Eye-specific active zone clustering underlies synaptic competition in the developing visual system

Author response:

The following is the authors’ response to the previous reviews

Reviewer #1 (Public review)

Summary:

This manuscript addresses the question of whether spontaneous activity contributes to the clustering of retinogeniculate synapses before eye opening. The authors re-analyze a previously published dataset to answer the question. The authors conclude that synaptic clustering is eye-specific and activity dependent during the first postnatal week. While there is useful information in this manuscript, I don't see how the data meaningfully supports the claims made about clustering.

In adult retinogeniculate connections, functionally specificity is supported by select pairings of retinal ganglion cells and thalamocortical cells forming dozens of synaptic connections in subcellular microcircuits called glomeruli. In this manuscript, the authors measure whether the frequency of nearby synapses is higher in the observed data than in a model where synapses are randomly distributed throughout the volume. Any real anatomical data will deviate from such a model. The interesting biological question is not whether a developmental state deviates from random. The interesting question is how much of the adult clustering occurs before eye opening. In trying to decode the analysis in this manuscript, I can't tell if the answer is 99% or 0.001%.

We thank the reviewer for their helpful critique through both rounds of review. We have refocused the manuscript on paired eye-specific measurements of active zone addition and spatial relationships among active zones at each age. All effect sizes and power values for each comparison are now reported in Table S2. These measures allow readers to gauge biological significance more transparently.

Strengths:

The source dataset is high resolution data showing the colocalization of multiple synaptic proteins across development. Added to this data is labeling that distinguishes axons from the right eye from axons from the left eye. The first order analysis of this data showing changes in synapse density and in the occurrence of multi-active zone synapses is useful information about the development of an important model system.

Weaknesses:

I don't think the analysis of clustering within this dataset improves our understanding of how the system works. It is possible that the result is clear to the authors based on looking at the images. As a reader trying to interpret the analysis, I ran into the following problems:

• It is not possible to estimate biologically meaningful effect sizes from the data provided. Spontaneous activity in the post natal week could be responsible for 99% or 0.001% of RGC synapse clustering.

• The sample size is too small for the kinds of comparisons being made. The authors point out that many STORM studies use an n of 1 while the authors have n = 3 for each of their six experimental groups. However, the critical bit is what kinds of questions you are trying to answer with a given sample size. This study depends on determining whether the differences between groups are due to age, genotype, or individual variation. This study also makes multiple comparisons of many different noisy parameters that test the same or similar hypothesis. In this context, it is unlikely that n = 3 sufficiently controls for individual variation.

We have revised the manuscript to focus on eye-specific differences, which are paired measurements collected at each age. We have measured effect sizes and performed power tests for all comparisons presented in the manuscript. These measurements are shown for every figure in a new supplemental table S2.

• There is no clear biological interpretation of the core measure of the publication, the normalized clustering index. The normalized clustering index starts with counting the fraction of single active zone synapses within various distances to the edge of synapses. This frequency is compared to a randomization model in which the positions of synapses are randomized throughout a volume. The authors found that the biggest deviation between the observed and randomized proximity frequency using a distance threshold of 1.5 um. They consider the deviation from the random model to be a sign of clustering. However, two RGC synapses 1.5 um apart have a good chance of coming from the same RGC axon. At this scale, real observations will, therefore, always look more clustered than a model where synapses are randomly placed in a volume. If you randomly place synapses on an axon, they will be much closer together than if you randomly place synapses within a volume. The authors normalize their clustering measure by dividing by the frequency of clustering in the normalized model. That makes the measure of clustering an ambiguous mix of synapse clustering, axon morphology, and synaptic density.

We have removed the “normalized clustering index”. “Clustered” inputs are now defined strictly as those that have a neighboring single active-zone (sAZ) synapse within 1.5 mm. For each type of input (sAZ and mAZ) we show 1) the ratio of clustered to isolated inputs for both eyes, and 2) the number of neighboring sAZs (Figure 4).

We agree with the reviewer that many synapses are likely made nearby along the same axon from an individual RGC. In this scenario, sAZ synapses that are nearby a neighboring mAZ input may be part of the same nascent bouton. And, sAZ synapses nearby other sAZ neighbors may ultimately mature into a mAZ input. At the same time, inputs from one RGC may form nearby other inputs from neighboring RGCs. We discuss these motifs and potential mechanisms of cell-autonomous and non-autonomous development (Lines 300-308).

• Other measures are also very derived. For instance, one argument is based on determining that the cumulative distribution of the distance of dominant-eye multi-active zone synapses with nearby single-active zone synapses from dominant-eye multi-active zone synapses is statistically different from the cumulative distribution of the distance of dominant-eye multi-active zones without nearby single-active zone synapses from dominant-eye multi-active zones. Multiple permutations of this measure are compared.

We have simplified the presentation to show all measured path lengths for every input. This allows the reader to see each of the inputs and their relative distances. We present these data for like-eye type interactions at P4 and P8 (Figures 5 and S5).

• There are major biological differences between groups that are difficult to control for. Between P2, P4, and P8, there are changes in cell morphology and synaptic density. There are also large differences in synapse density between wild type and KO mice. It is difficult to be confident that these differences are not responsible for the relatively subtle changes in clustering indices.

• Many claims are based on complicated comparisons between groups rather than the predominating effects within the data. It is noted that: "In KO mice, dominant eye projections showed increased clustering around mAZ synapses compared to sAC synapses suggesting partial maintenance of synaptic clustering despite retinal wave defects". In contrast, I did not notice any discussion of the fact that the most striking trend in those measures is that the clustering index decreases from P2 to P8.

Related to the points above, we have revised the manuscript to focus on eye-specific release site addition and spatial relationships. For clarity, we have removed the clustering index and instead present ratios of clustered and isolated inputs, the number of sAZ synapses near each input type, and distance between like-eye mAZ inputs (Figure 4).

• Statistics are improperly applied. In my first review I tried to push the authors to calculate confidence intervals for two reasons. First, I believed the reader should be able to answer questions such as whether 99% or 0.01% of RGC synaptic clustering occurred in the first postnatal week. Second, I wanted the authors to deal with the fact that n=3 is underpowered for many of the questions they were asking. While many confidence intervals can now be found leading up to a claim, it is difficult to find claims that are directly supported by the correct confidence interval. Many claims are still incorrectly based on which combinations of comparisons produced statistically significant differences and which combinations did not.

We have substantially revised the manuscript to focus on within-group paired effects between eye-of-origin. We performed power tests for all statistical presentations and effect sizes and powers are presented for every figure in a new supplemental table S2. To simplify the manuscript and make it easier to read, we report confidence interval measurements in a separate supplemental table S3.

Reviewer #2 (Public review):

Summary:

This study provides a valuable data set showing changes in the spatial organization of synaptic proteins at the retinogeniculate connection during a developmental period of active axonal and synaptic remodeling. The data collected by STORM microscopy is state-of-the-art in terms of the high-resolution view of the presynaptic components of a plastic synapse. The revision has addressed many, but not all, of the initial concerns about the authors interpretation of their data. However, with the revisions, the manuscript has become very dense and difficult to follow.

We greatly appreciate the reviewer’s thoughtful comments through two rounds of review. To improve the clarity of the manuscript, we have substantially revised the work to streamline the narrative, clearly define terminology, and simplify data presentations, allowing readers to more directly interpret results and their implications.

Strengths:

The data presented is of good quality and provides an unprecedented view at high resolution of the presynaptic components of the retinogeniculate synapse during active developmental remodeling. This approach offers an advance to the previous mouse EM studies of this synapse because the CTB label allows identification of the eye from which the presynaptic terminal arises.

Weaknesses:

From these data the authors conclude that eye-specific increase in mAZ synapse density occur over retinogeniculate refinement, that sAZ synapses cluster close to mAZ synapses over age, and that this process depends on spontaneous activity and proximity to eye-specific mAZ synapses. While the interpretation of this data set is much more grounded in this revised submission, some of the authors' conclusions/statements still lack convincing supporting evidence.

This includes:

(1) The conclusion that multi-active zone synapses are loci for synaptic clustering. This statement, or similar ones (e.g., line 407) suggest that mAZ synapses actively or through some indirect way influence the clustering of sAZ synapses. There is no evidence for this. Clustering of retinal synapses are in part due to the fact that retinal inputs synapse on the proximal dendrites. With increased synaptogenesis, there will be increased density of retinal terminals that are closely localized. And with development, perhaps sAZ synapses mature into mAZ synapses. This scenario could also explain a large part of this data set.

We thank the reviewer for their comment. We have removed the ambiguous phrasing and clarified the manuscript to explicitly discuss alternative interpretations consistent with the results (Lines 300-308). This includes a discussion of sAZ synapse maturation into mAZ inputs (Lines 294-296).

(2) The conclusion that, "clustering depends on spontaneous retinal activity" could be misleading to the reader given that the authors acknowledge that their data is most consistent with a failure of synaptogenesis in the mutant mice (in the rebuttal). Additionally clustering does occur in CTB+ projections around mAZ synapses.

We have removed the highlighted phrase and revised the manuscript to focus on differences in release site addition between eye-of-origin. We clarified our discussion of activity-dependent changes to state that synapses fail to form in the mutant and synaptic clustering was reduced (Lines 324-330).

(3) Line 403: "Since mAZ synapses are expected to have a higher release probability, they likely play an important role in driving plasticity mechanisms reliant on neurotransmission.":What evidence do the authors have that mAZ are expected to have higher release probability?

We thank the reviewer for their careful reading. Because they have several active zones, mAZ synapses are expected to have a higher number of release sites (N), which could be independent of release probability at any individual active zone (Pr). We have removed the reference to release probability. Instead, we maintain focus on active zone number.

Reviewer #3 (Public review):

This study is a follow-up to a recent study of synaptic development based on a powerful data set that combines anterograde labeling, immunofluorescence labeling of synaptic proteins, and STORM imaging (Cell Reports, 2023). Specifically, they use anti-Vglut2 label to determine the size of the presynaptic structure (which they describe as the vesicle pool size), anti-Bassoon to label active zones with the resolution to count them, and anti-Homer to identify postsynaptic densities. Their previous study compared the detailed synaptic structure across the development of synapses made with contra-projecting vs. ipsi-projecting RGCs and compared this developmental profile with a mouse model with reduced retinal waves. In this study, they produce a new detailed analysis on the same data set in which they classify synapses into "multi-active zone" vs. "single-active zone" synapses and assess the number and spacing of these synapses. The authors use measurements to make conclusions about the role of retinal waves in the generation of same-eye synaptic clusters, providing key insight into how neural activity drives synapse maturation.

Strengths:

This is a fantastic data set for describing the structural details of synapse development in a part of the brain undergoing activity-dependent synaptic rearrangements. The fact that they can differentiate eye of origin is what makes this data set unique over previous structural work. The addition of example images from EM data set provides confidence in their categorization scheme.

Weaknesses:

Though the descriptions of synaptic clusters are important and represent a significant advance, the authors conclusions regarding the biological processes driving these clusters are not testable by such a small sample. This limitation is expected given the massive effort that goes into generating this data set. Of course the authors are free to speculate, but many of the conclusions of the paper are not statistically supported.

We thank the reviewer for their helpful comments throughout the revision process. We have substantially modified the manuscript to reframe the work around release site addition during eye-specific competition. Power tests and effect size measurements are presented for every figure in a new supplemental table S2.

Reviewer #2 (Recommendations for the authors):

(1) Authors should discuss that it is not clear what the relationship is between sAZ and mAZ, and sAZ could turn into a mAZ. This is not unreasonable that the number of AZ/bouton increases with development given that in the adult rodent retinogeniculate bouton, there is an average of 27 active zones (Budisantoso et al, 2012).

We thank the reviewer for their helpful suggestion. We have added a discussion of the relationship between sAZ and mAZ inputs and the point that sAZ synapses may mature into mAZ synapses (Lines 294-296). We now reference the work of Budisantoso et al., J. Neurosci. 2012.

(2) The authors should clarify how the statistics are calculated for the normalized clustering index (figure 3B, C). For ratios of values each with variance, the variance is summed when calculating SEM.

For clarity, we have removed the normalized clustering index analysis. We have simplified the work to present a clear definition of clustered and unclustered inputs, where clustering is defined by the presence of a nearby neighboring synapse within 1.5mm. We present the ratio of clustered and unclustered inputs for each input type and eye-of-origin. We also show the number of sAZ synapses nearby each clustered input (Figure 4).

(3) The authors have significantly clarified the terminology that they use in the text. This is much appreciated. However, it would be helpful to the naïve reader if they could define their use of the word "synapse" as referring to individual active zones/release sites or to terminals/boutons. For example:

Line 378: "Prior electron microscopy studies in the mouse found limited evidence of convergent synaptic clustering from neighboring RGCs at postnatal day 8 (10, 13), suggesting that the mAZ synapses seen in STORM images are single retinogeniculate terminals. The lack of synaptic convergence in prior EM reconstructions at P8 implies that early clustering around mAZ synapses may result from local output clustering within individual RGC arbors.":

What do the authors mean by "convergent synaptic clustering": do they mean clustering of release sites from different RGC inputs? And what does "local output clustering" mean?

We thank the reviewer for their suggestion to use clear terminology. We have revised the manuscript to define our use of the term “synapse” as a single active zone/release site (Lines 134-136). We refer to mAZ boutons in STORM data as “inputs”. We have revised the discussion of prior EM studies (Lines 130-132) and clarified all discussions of synaptic clustering throughout the work.

(4) While the authors argue that the retina-specific β2-nAChR mice exhibit disrupted retinal waves and defects in eye specific segregation, the authors are studying issues of active zone density which may depend on mechanisms depending on the postsynaptic neuron. This should be acknowledged.

We have updated the text to discuss the fact that postsynaptic mechanisms are also critical for the refinement of eye-specific synapses (Lines 332-340). We have added several additional references to the manuscript accordingly.

Reviewer #3 (Recommendations for the authors):

The authors have addressed many of my original concerns. The additional description of criteria for categorizing synapses, showing all the data points, gives the reader a stronger sense of where the numbers in the quantification come from. Replacing the "complex/simple" distinction with the "multi/single active zone" and the other clarifying text was effective. The addition of the EM data was also a very nice example to help interpret STORM images. It does appear there was no quantification on this EM data set and perhaps just a few example images were taken as "proof of principle". If, by chance, the authors have more EM images to make a data set of them that allows for some quantification, that would be great to add.

We thank the reviewer for their helpful comments on the manuscript through both rounds of review. The EM data we collected were 2D images of a subset of physical sections at postnatal day 8. Most dAPEX2(+) profiles had a single active zone, but a definitive identification would require 3D imaging so that each terminal can be assessed in its entirety for release sites that might be missed in a single cross section. Similarly, multi-active zone boutons are positively identified in 2D images, but definitive measurements of AZ number would require 3D information. We analyzed our 2D EM images and present a plot of dAPEX2(+) profile size versus active zone number below. These measures are positively correlated (r = 0.74), with larger profiles containing more active zones.

Unfortunately, we are not currently equipped to perform volumetric EM imaging at our home institution and are concerned that analysis of 2D data may be inconclusive. For these reasons, we are opting to maintain a qualitative presentation of our current EM results and we look forward to collaborating with other experts to achieve volumetric EM reconstructions in the future

Changes from the earlier version to the revised version

Abstract

Abstract

Added in the revised version: Spatially clustered synaptic inputs enable local dendritic computations important for learning, memory, and sensory processing. In the mammalian visual system, individual retinal ganglion cell (RGC) axons form clustered terminal boutons containing multiple active zones onto relay cell dendrites in the dorsal lateral geniculate nucleus (dLGN). This mature architecture arises through the addition of release sites, which strengthens selected afferents while weaker inputs are pruned. Following eye-opening, spontaneous activity and visual experience promote synaptic refinement and bouton clustering after binocular inputs have segregated. However, anatomical changes in release site addition and spatial patterning during earlier stages of eye-specific competition are not well understood. To investigate this, we examined the spatial organization of eye-specific active zones in wild type mice and a mutant line with disrupted cholinergic retinal waves. Using volumetric super-resolution single-molecule localization microscopy and electron microscopy, we found that individual retinogeniculate boutons begin forming multiple nearby presynaptic active zones during the first postnatal week. Both eyes generate these “multi-active-zone” (mAZ) inputs throughout refinement, but the dominant-eye forms more numerous mAZ contacts, each with more active zones and larger vesicle pools. At the height of competition, the non-dominant-eye projection adds many single active zone (sAZ) synapses. Mutants with abnormal cholinergic retinal waves still form mAZ inputs, but develop fewer sAZ synapses and show reduced synapse clustering in projections from both eyes. Together, these findings reveal activity-dependent, eye-specific differences in release site addition during synaptic competition in circuits essential for visual perception and behavior.

Removed in the revised version: Co-active synaptic connections are often spatially clustered to facilitate local dendritic computations underlying learning, memory, and basic sensory processing. In the mammalian visual system, retinal ganglion cell (RGC) axons converge to form clustered synaptic inputs that enable local signal integration in the dorsal lateral geniculate nucleus (dLGN) of the thalamus. While visual experience promotes retinogeniculate synapse clustering after eye-opening, the earliest events in cluster formation prior to visual experience are unknown. Here, using volumetric super-resolution single-molecule localization microscopy and eye-specific labeling of developing retinogeniculate synapses in mice, we show that synaptic clustering is eye-specific and activity-dependent during retinogeniculate refinement in the first postnatal week. We identified a subset of retinogeniculate synapses with multiple active zones that are surrounded by like-eye synapses and depleted of synapse clustering from the opposite eye. In mutant mice with disrupted spontaneous retinal wave activity, synapses with multiple active zones still form, but do not exhibit the synaptic clustering seen in controls. These results highlight a role for spontaneous retinal activity in regulating eye-specific synaptic clustering in circuits essential for visual perception and behavior.

Introduction

Earlier version: A key mechanism for neuronal signal processing is the formation of spatially clustered synaptic inputs, which facilitate local computations within individual dendrites (1-3). Through biophysical signal integration mechanisms, these synaptic clusters play critical computational roles in learning, memory, and sensory processing underlying cognition and behavior (4-7). During circuit development, the formation of synaptic clusters is regulated by both spontaneous and sensory-driven neural activity, which helps to stabilize or eliminate individual synapses and establish mature connectivity patterns (5, 6, 8).

Revised version: A key mechanism for neuronal signal processing is the formation of spatially clustered synapses that facilitate local computations within individual dendrites 1–3. Through biophysical signal integration mechanisms, neighboring synaptic inputs play critical computational roles in learning, memory, and sensory processing underlying cognition and behavior 4–7. During circuit development, both spontaneous and sensory-driven neural activity regulate synaptic clustering, stabilizing some synapses and eliminating others to establish mature connectivity patterns 6–8.

Added in the revised version: The refinement of retinal inputs to the dorsal lateral geniculate nucleus (dLGN) of the thalamus is a model example of activity-dependent synaptic clustering 9. Over development, boutons within the terminal arbors of individual retinogeniculate axons cluster progressively 10–13, creating multiple synaptic contacts onto individual postsynaptic targets 14–16. Concurrently, developing RGC boutons add presynaptic release sites such that some mature retinogeniculate terminals contain several dozen active zones (AZs) 16, 17. The formation of clustered boutons with multiple release sites subserves a critical “driver” function of retinogeniculate input 18, 19.

Added in the revised version: Retinogeniculate bouton clustering depends on visual experience, and dark-rearing after eye-opening reduces clustering within individual axon arbors 12. However, less is known about synaptic spatial relationships and activity-dependent release site refinement during eye-specific competition before eye-opening. During early axon ingrowth, individual RGCs extend sparse side branches in the incorrect eye-specific territory 20–22. These transient branches form synapses that contain fewer presynaptic vesicles than those made by the same axon in the correct eye-specific domain 23.

Added in the revised version: Functional recordings after eye-opening show that while many RGC inputs are pruned, the remaining inputs strengthen by adding release sites 24. This leads to glutamate spillover and cross talk between RGC inputs that increases postsynaptic relay neuron excitability 25. Thus, early eye-specific differences in release site number or spacing may contribute to binocular input refinement prior to eye-opening.

Added in the revised version: In a previous study from our laboratory, we used volumetric stochastic optical reconstruction microscopy (STORM), anterograde tract tracing, and immunohistochemical labeling of synaptic proteins to show that dominant-eye synapses contain more total vesicles and have more vesicles near the AZ (putative docking) compared with non-dominant-eye synapses 26. However, we did not examine whether eye-specific inputs differ in AZ number or AZ spatial proximity (clustering) during synaptic competition. To address this gap, we reanalyzed our published dataset, focusing on eye-specific AZ spatial relationships. We found that projections from each eye form a subset of retinogeniculate inputs in which several (2-4) AZs share a common cluster of presynaptic vesicles. We refer to these synapses as multi-active-zone (mAZ) inputs. All other retinogeniculate inputs contained one single active zone and its associated vesicle cluster, and we term these single-active-zone (sAZ) synapses.

Added in the revised version: During eye-specific synaptic competition, the dominant-eye projection formed more mAZ inputs, each with more AZs and a larger presynaptic vesicle pool compared to the non-dominant eye projection. Similarly, the dominant-eye had higher vesicle signal at sAZ inputs. At the peak of synaptic competition midway through the first postnatal week, the non-dominant-eye formed numerous sAZ inputs, equalizing the global synapse density between the two eyes. These eye-specific AZ patterns were disrupted in a mutant mouse line with abnormal stage II cholinergic retinal waves and retinogeniculate segregation defects. Our results demonstrate that eye-specific presynaptic release site clustering arises before eye-opening and is impacted by defects in spontaneous retinal activity.

Removed in the revised version: A model example of activity-dependent synaptic cluster formation is the refinement of retinal inputs to the dorsal lateral geniculate nucleus (dLGN) of the thalamus (9). Electron microscopy (EM) reconstructions of the mouse dLGN show that retinal ganglion cell (RGC) inputs converge to form “complex” synaptic clusters known as glomeruli (10-14). Each glomerulus contains multiple RGC axon terminal boutons formed onto a dendritic branch of a dLGN relay neuron. These boutons vary from small terminals with a single active zone (AZ) to large or perforated synapses containing multiple AZs (10-13). Ultrastructural evidence of synaptic integration is supported by experiments using transsynaptic (15) and Brainbow-based RGC labeling (11, 13), optogenetic stimulation of RGC axons (16, 17), and calcium imaging of retinogeniculate boutons (18) confirming the clustering of RGC inputs to relay neuron dendrites in the dLGN. Because individual glomeruli often receive inputs from multiple RGCs that encode either similar or distinct visual features (18), proper developmental wiring of RGC bouton clusters is critical for local dendritic integration that drives visual responses in the adult brain.

Removed in the revised version: Previous studies have shown that synaptic cluster development in neural circuits depends on spatiotemporally correlated synaptic activity to regulate synaptogenesis, pruning, and plasticity (6, 19). In the developing retinogeniculate system, there are two sources of correlated activity that may contribute to synaptic clustering: 1) visual experience that drives topographic activation of neighboring RGCs and 2) spontaneous retinal wave activity that correlates burst firing of neighboring RGCs prior to eye-opening. Consistent with experience-dependent plasticity, retinogeniculate bouton clustering increases after eye-opening (10, 13, 20) and visual deprivation reduces clustering (20). However, these experience-dependent synaptic clustering effects occur after eye-opening, when eye-specific retinogeniculate segregation is complete (21-23). Before photoreceptor-mediated visual onset, retinal waves generate spatiotemporal correlations in RGC burst activity that are expected to facilitate Hebbian strengthening of co-active synapses (24). Whether this correlated activity promotes retinogeniculate synapse clustering prior to eye-opening is unknown.

Removed in the revised version: To address this question, we used volumetric super-resolution microscopy, anterograde tract tracing, and immunohistochemical synaptic protein labeling to investigate the local spatial relationships between eye-specific synapses in the dLGN prior to photoreceptor-mediated visual experience. During eye-specific segregation (postnatal days 2-8), we identified a subset of retinogeniculate inputs that formed multiple active zones (mAZ synapses), while all other retinal input synapses contained only a single active zone (sAZ synapses). Comparing the spatial relationships between all retinogeniculate synapses, we observed non-random, eye-specific clustering of sAZ synapses around mAZ synapses. During eye-specific synaptic competition, both the “winning” and “losing” eye-specific inputs showed increased sAZ clustering when like-eye mAZ synapses were in close proximity. In contrast, mAZ synapses from the “losing” eye projection were less clustered when located near a mAZ synapse from the “winning” eye.

Removed in the revised version: These patterns of eye-specific synaptic clustering were absent in a genetic mutant mouse line with abnormal stage II cholinergic retinal waves and retinogeniculate segregation defects. These results demonstrate that spontaneous retinal activity regulates retinogeniculate synaptic clustering prior to eye-opening and are consistent with the action of non-cell-autonomous synaptic stabilization and punishment signals underlying eye-specific competition in the developing visual system.

Results

Earlier version: Retinogeniculate synapses form multiple active zones during eye-specific competition

Revised version: Retinogeniculate inputs form multiple active zones during eye-specific competition

Removed in the revised version: To investigate synaptic clustering during eye-specific segregation, we reanalyzed a previously published volumetric super-resolution imaging dataset from our laboratory (25). Using volumetric STochastic Optical Reconstruction Microscopy (STORM) (26), we collected super-resolution imaging data from the dLGNs of wild-type (WT) mice at three postnatal time points (P2, P4, and P8) (Figure 1A). We labeled eye-specific synapses by monocular injection of Alexa Fluor-conjugated cholera toxin subunit B tracer (CTB) together with immunostaining for presynaptic Bassoon, postsynaptic

Removed in the revised version: Retinogeniculate synapses form multiple active zones during eye-specific competition.

Removed in the revised version: (A) Experimental design: CTB-Alexa 488 was injected into the right eye of wild-type (WT) and β2KO mice. One day after the treatment, tissue was collected from the left dLGN at P2, P4, and P8. The red squares indicate the STORM imaging regions. (B) Representative multi-active zone (mAZ, left panels) and single-active zone (sAZ, right panels) synapses in WT (top panels) and β2KO mice (bottom panels) at each age. Arrowheads point to individual Bassoon clusters (active zones) within each synapse. (C) Representative images of mAZ retinogeniculate synapses in the mouse dLGN at P8. Electron microscopy images show darkly-stained dAPEX2(+) mitochondria within RGC boutons. Arrowheads point to electron-dense material at the postsynaptic density (PSD) apposed to individual active zones with clustered presynaptic synaptic vesicles. Scale bars are 1 µm in all images. (D) Representative CTB(+) dominant-eye (top panels) and CTB(-) non-dominant-eye (bottom panels) mAZ synapses in a WT P8 sample, showing synaptic (left panels), CTB (middle panels), and merged immunolabels (right panels). (E and F) Eye-specific mAZ synapse density (E) and mAZ synapse fraction (F) across development in WT (top panel) and β2KO mice (bottom panel). Black dots represent mean values from separate biological replicates and black lines connect measurements within each replicate (N=3 for each age and genotype). Error bars represent group means ± SEMs. Statistical significance between CTB(+) and CTB(-) synapse measurements was assessed using paired T-tests. *: p<0.05, **: p<0.01, ***: p<0.001.

Earlier version: Homer1, and presynaptic vesicular glutamate transporter 2 (VGluT2) proteins (25). We collected separate image volumes (∼45K μm3 each) from three biological replicates at each developmental time point. To assess the impact of spontaneous retinal activity on synaptic development across the same time period, we performed identical experiments in a knockout mouse line lacking the beta 2 subunit of the nicotinic acetylcholine receptor (β2KO), which disrupts spontaneous cholinergic retinal wave activity, eye-specific axonal segregation, and retinogeniculate synapse development (21, 25, 27-35). Because eye-specific segregation is incomplete until ∼P8, we limited our analysis to the future contralateral eye-specific region of the dLGN, which is reliably identified across all stages of postnatal development (Figure 1A, see also Materials and Methods).

Revised version: To investigate active zone refinement during eye-specific segregation, we reanalyzed a volumetric super-resolution imaging dataset previously published by our laboratory 26. We used volumetric STochastic Optical Reconstruction Microscopy (STORM) 27 to image the dLGNs of wild-type (WT) mice at three postnatal ages (P2, P4, and P8) (Fig. 1A). We labeled eye-specific inputs by monocular injection of Alexa Fluor-conjugated cholera toxin subunit B tracer (CTB) together with immunostaining for presynaptic Bassoon, postsynaptic Homer1, and presynaptic vesicular glutamate transporter 2 (VGluT2) proteins (Fig. 1B). We collected separate image volumes (∼45K μm3 each) from three biological replicates at each age. To assess the impact of spontaneous retinal activity on synaptic development across the same time period, we performed identical experiments in β2-knockout (β2KO) mice lacking the beta 2 subunit of the nicotinic acetylcholine receptor, a mutation that disrupts spontaneous cholinergic retinal wave activity, eye-specific segregation, and retinogeniculate synapse development 22, 26, 28–36 (Fig. 1A). Because eye-specific segregation is incomplete until ∼P8, we limited our re-analysis to the future contralateral eye-specific region of the dLGN, which is reliably identified across all stages of postnatal development (Fig. 1A, see also Materials and Methods).

Added in the revised version: Retinogeniculate boutons form multiple active zones during eye-specific competition.

Added in the revised version: (A) Experimental design. CTB-Alexa 488 was injected into the right eye of wild-type and β2KO mice. One day after the treatment, tissue was collected from the left dLGN at P2, P4, and P8. Red squares indicate the STORM imaging regions that were analyzed. (B) Representative examples of individual sAZ and mAZ inputs, with corresponding active zone counts ranging from one to three. Upper panels show Z-projections of inputs and lower panels show the corresponding 3D volume. Arrowheads point to individual Bassoon clusters (active zones) paired with postsynaptic Homer1 labels within each input. All examples are from a WT P8 sample. (C) Electron micrographs of mAZ retinogeniculate inputs in a P8 ET33Cre::ROSA26LSL-Matrix-dAPEX2 mouse. Darkly-stained dAPEX2(+) mitochondria are present within ipsilaterally-projecting RGC terminals. Arrowheads point to electron-dense material at the postsynaptic density, apposed to individual active zones with clustered presynaptic synaptic vesicles.

Added in the revised version: By analyzing synapses in the contralateral dLGN from eighteen mice across three ages and two genotypes (Table S1), STORM revealed two classes of retinogeniculate inputs distinguished by active zone (AZ) number (Fig. 1B). We defined each retinogeniculate input as a single contiguous VGluT2 cluster together with all its associated presynaptic (Bassoon) and postsynaptic (Homer1) paired synaptic labels. Using this definition, inputs that had multiple (2-4) Bassoon AZs were classified as multi-active-zone (mAZ) inputs, while those with a single Bassoon AZ were designated single-active-zone (sAZ) inputs (Fig. 1B). Most mAZ inputs contained two AZs (∼70-90%, varying with age, genotype, and eye-of-origin); smaller proportions contained three AZs (∼10-20%) or four or more AZs (<5%) (Table S1).

Added in the revised version: Each mAZ input could be a single terminal bouton with several AZs, or a cluster of sAZ synapses within separate boutons 10, 11, 15, 16, 37. To address this, we used electron microscopy (EM) to image retinogeniculate terminals in the dLGN at P8. We generated a transgenic line expressing mitochondrial matrix-targeted dimeric dAPEX2 reporter 38 in ipsilaterally-projecting RGCs 39–41, providing unambiguous mitochondrial labeling in ipsiRGC axons. EM images confirmed the presence of individual retinogeniculate boutons with multiple active zones, consistent with our STORM data (Fig. 1C). Previous EM reconstructions of retinogeniculate inputs reported no evidence of RGC bouton convergence at the end of the first postnatal week 11. Together, these results suggest that mAZ inputs in STORM images are single RGC terminals that house several closely spaced release sites. Hereafter, we use the word “synapse” only to refer to each partnered active zone (Bassoon/Homer1 pairs); mAZ inputs contain several synapses, while each sAZ input is one synapse.

Added in the revised version: Changes in eye-specific input density during synaptic competition

Removed in the revised version: Across our dataset containing tens of thousands of identified synapses (Supplemental Table 1) collected from different ages and genotypes, STORM images revealed two subtypes of retinogeniculate synapses distinguished by the number of active zones (AZs) in the presynaptic terminal (Figure 1B). One subtype included synapses with multiple (2-4) Bassoon(+) AZs (hereafter referred to mAZ synapses), while the other consisted of retinogeniculate synapses with a single Bassoon(+) AZ (sAZ synapses) (Figures 1B and S1A). Each mAZ synapse could be a single terminal, or, alternatively, several clustered RGC boutons each of which is a sAZ synapse (10-13). To examine the ultrastructure of mAZ synapses, we used electron microscopy (EM) to image retinogeniculate boutons in the dLGN. We generated a transgenic mouse line expressing the mitochondrial matrix-targeted dimeric dAPEX2 reporter (36) in ipsilaterally projecting RGCs (37) to provide unambiguous mitochondrial labeling in RGC axons and synaptic terminals. EM images confirmed the presence of large retinogeniculate terminals with multiple active zones, consistent with our STORM data (Figure 1C).

Earlier version: Because our previous analysis of retinogeniculate synaptic refinement reported global changes in eye-specific synapse density regardless of AZ number (25), we first sought to separate mAZ and sAZ synapses and determine their relative densities and ratios during eye-specific segregation. To determine the eye-of-origin for each retinogeniculate synapse, we measured the colocalization of CTB signal with VGluT2 (Figure 1D). By imaging in the contralateral eye-specific region relative to the CTB-injected eye, we defined CTB(+) VGluT2 clusters as “dominant-eye” synapses. Conversely, CTB(-) VGluT2 clusters were classified as “non-dominant eye” synapses. Our previous analysis of binocular CTB control injection data confirmed the high efficiency of retinogeniculate synapse labeling by anterograde tracing (>97% of synapses labeled at P4 and P8), thereby enabling accurate assignment of eye-specific synapses in the mouse brain (25).

Revised version: In our previous analysis, we reported global eye-specific synapse densities that reflected the combined mAZ and sAZ inputs. Dominant-eye synapse density was greater than that of the non-dominant eye at P2 and P8, but eye-specific synapse densities were equivalent at P4 during the peak of synaptic competition 26. To determine how these two input classes evolve during competition, we quantified the densities and percentages of mAZ and sAZ inputs in WT and β2KO mice (Fig. 2A). Eye-of-origin for each retinogeniculate input was assigned by colocalizing CTB signal with VGluT2 (Fig. 2B). Binocular CTB control injections showed that anterograde tracing labeled >97% of retinogeniculate synapses at P4 and P8, ensuring accurate eye-specific assignment 26. Within the contralateral eye-specific region of the dLGN, CTB(+) VGluT2 clusters were classified as “dominant-eye” inputs and CTB(-) VGluT2 clusters as “non-dominant eye” inputs (Fig. 2B).

Added in the revised version: Changes in eye-specific input density during synaptic competition.

Removed in the revised version: After assigning eye-of-origin to all retinogeniculate synapses, we observed that the density of CTB(+) mAZ synapses was significantly higher than CTB(-) mAZ synapses at every age (Figure 1E). In contrast, sAZ density showed eye-specific differences at P2 and P8, but not at P4 (Figure S1B, 5/95% difference confidence interval -0.021∼0.008 synapses / μm3). By adding eye-specific mAZ and sAZ numbers, we found that at P4 the total synapse density is equivalent between the two eyes (5/95% difference confidence interval -0.014∼0.011 synapses / μm3) and this effect is driven by an increase in the number of CTB(-) ipsilateral sAZ synapses at P4 (Figure S1B). This effect was not seen in β2KO mice, which showed eye-specific differences in the density of both mAZ and sAZ synapses at P4 (Figure 1E and S1B, 5/95% difference confidence interval 0.004∼0.010 mAZ synapses / μm3 and 0.007∼0.020 sAZ synapses / μm3 at P4). In WT mice, the fraction of dominant-eye mAZ synapses was ∼20% at P2 and ∼29% at P8 while non-dominant-eye mAZ synapses comprised∼14% and ∼17% at the same time points (Figure 1F, upper panel). In β2KO mice, dominant-eye mAZ synapses represented 17/23% (P2/8) and non-dominant eye mAZ synapses accounted for 10/14% (P2/P8) of the total population. These results indicate that a population of mAZ synapses still forms in β2KO mice despite an overall reduction in synapse density when spontaneous retinal activity is disrupted (25).

Removed in the revised version: Multi-active zone synapses undergo eye-specific vesicle pool maturation

Removed in the revised version: In our previous study of retinogeniculate refinement, we reported eye-specific differences in VGluT2 volume underlying synaptic competition, but did not investigate the relative maturation of presynaptic vesicle pools in sAZ and mAZ synapses (25). To determine these differences, we measured the total VGluT2 volume for mAZ and sAZ synapses. In WT mice, both mAZ (Figure 2A, left panel) and sAZ (Figure 2B, left panel) synapses showed eye-specific differences in presynaptic VGluT2 volume at each time point. During eye-specific competition at P4 in WT mice, the median dominant-eye VGluT2 cluster volume was 372% larger (marginal mean with 95% confidence interval on a logarithmic scale: 0.551 ± 0.294 log10[μm3]) compared to non-dominant-eye synapses (Figure 2A, left panel). In contrast, β2KO mice showed a smaller difference (110%, marginal mean with 95% confidence interval on a logarithmic scale: 0.279 ± 0.107 log10[μm3]) between eye-specific mAZ synapse VGluT2 volumes at the same time point (Figure 2A, right panel). For sAZ synapses at P4, the magnitudes of eye-specific differences in VGluT2 volume were significantly smaller: 135% in WT, (Figure 2B, left panel, marginal mean with 95% confidence interval of CTB-specific difference on a logarithmic scale: 0.337 ± 0.275 log10[μm3], the difference of CTB-specific contrast between mAZ and sAZ synapses: 0.214 ± 0.102 log10[μm3]) and 41% inβ2KO) (Figure 2B, right panel, marginal mean with 95% confidence interval of CTB-specific difference on a logarithmic scale: 0.167 ± 0.119 log10[μm3], the difference of CTB-specific contrast between mAZ and sAZ synapses: 0.112 ± 0.040 log10[μm3]). These data show that the eye-specific pool size differences in mAZ synapses are 2.76-fold greater in WT and 2.68-fold greater in β2KO compared to sAZ synapses, indicating that synapse type-specific differences persist despite defects in retinal waves.

Removed in the revised version: Multi-active zone synapses undergo eye-specific vesicle pool maturation.

Removed in the revised version: (A) Violin plots showing the distribution of VGluT2 cluster volume for mAZ synapses in WT and β2KO mice at each age. The width of each violin plot reflects the relative synapse proportions at each volume across the entire grouped data set (N=3 biological replicates). The maximum width of the violin plots was normalized across mAZ and sAZ groups. The black dots represent the median value of each biological replicate (N=3), and the black horizontal lines represent the median value of all synapses grouped across replicates. Black lines connect measurements of CTB(+) and CTB(-) populations from the same biological replicate. Statistical significance was determined using a mixed model ANOVA. Black asterisks indicate significant eye-specific differences at each developmental time point. (B) Violin plots similar to (A) show the distributions of VGluT2 cluster volume for sAZ synapses in WT and β2KO mice at each age. (C) Average number of AZs (individual bassoon clusters) per mAZ synapse in WT (top panel) and β2KO mice (bottom panel). Black dots represent mean values from separate biological replicates and black lines connect measurements within each replicate. (D and E) Total (D) and average/AZ (E) VGluT2 cluster volume as a function of AZ number for all synapses in WT P4 samples (top panel) and β2KO P4 samples (bottom panel). Black dots represent mean values from separate biological replicates and black lines connect measurements within each replicate. A table of the total eye-specific synapse numbers of each type (mAZ and sAZ) and AZ number for each biological replicate in both genotypes is included in the Supplemental Materials (Table S1). In (C-E), error bars represent means ± SEMs. Statistical significance was assessed using paired T-tests. In all panels, *: p<0.05, **: p<0.01, ***: p<0.001. N=3 biological replicates for each age and genotype.

Removed in the revised version: To determine if the presynaptic vesicle pool volume scales with the number of active zones, we compared the number of Bassoon clusters toVGluT2 cluster volume for each mAZ synapse. In both WT (Figure 2C, top panel) and β2KO mice (Figure 2C, bottom panel), mAZ synapses were associated with an average of 2-3 Bassoon clusters in the first postnatal week. In WT mice, the number of AZs in CTB(+) dominant-eye mAZ synapses was significantly higher than in non-dominant CTB(-) mAZ synapses at P4 and P8 (Figure 2C, top panel). However, this difference was not observed at P4 in β2KO mice (Figure 2C, bottom panel. 5/95% confidence interval -0.0005∼0.015). For CTB(+) dominant-eye synapses in both WT and β2KO mice, vesicle pool size was positively correlated with AZ number (Figure 2D, Figure S2A/B; correlation coefficients: WT [0.99] and β2KO [0.95]). This correlation was also present for CTB(-) non-dominant-eye synapses, but with a smaller magnitude at P4 and P8 (correlation coefficients: WT [0.97] and β2KO [0.90]. Dividing the total presynaptic VGluT2 volume by the AZ number revealed a consistent vesicle volume per AZ for both sAZ and mAZ synapses(Figure 2E, S2C/D). In dominant-eye CTB(+) synapses, 5/95% confidence intervals of mAZ VGluT2 volume were -61%∼50% of sAZ synapses in WT mice and -32%∼11% in β2KO mice. In non-dominant-eye CTB(-) synapses, 5/95% confidence intervals of mAZ VGluT2 volume were -88%∼51% of sAZ synapses in WT mice and -43%∼5% in β2KO mice. The results indicate that for each synapse, vesicle pool volume increases linearly with the number of AZs.

Removed in the revised version: Multi-active zone synapses are loci for synaptic clustering

Removed in the revised version: Eye-specific competition is thought to result from the stabilization of co-active dominant-eye RGC inputs and the elimination of non-dominant-eye inputs that are out of sync with neighboring synapses (38-40). In this context, mAZ synapses with greater release probability may play an important role in the local refinement of nearby synapses (41-43). To determine whether retinogeniculate synapses cluster prior to eye-opening, we measured the proximity of eye-specific sAZ synapses to mAZ synapses and other sAZs synapses. We then compared these distances to a randomized synapse distribution, where sAZ synapse positions were shuffled within the neuropil volume while mAZ positions remained fixed. For every sAZ and mAZ synapse in the dataset, we identified neighboring sAZ synapses as “nearby” when their weighted centroid was within a defined distance from the edge of the target center synapse (Figure 3A, a mAZ synapse in this case). We tested search radii between 1-4 μm around both CTB(+) and CTB(-) mAZ synapses, measuring the ratio of nearby sAZ synapses in the original versus randomized data sets (Figures S3A/B). The largest differences between the original and randomized distributions occurred within 1-2 μm search radius, suggesting non-random clustering of sAZ synapses near mAZ synapses from the same eye (Figure S3A/B). We did not observe this clustering pattern when sAZ synapses and mAZ synapses originated from different eyes (Figure S3C/D). This analysis defined “nearby” sAZ synapses as those within 1.5 μm of a mAZ synapse edge (Figure 3A, S3).

Removed in the revised version: Multi-active zone synapses are loci for synaptic clustering.

Earlier version: (A) A representative mAZ synapse (arrow) in a WT P8 dLGN with nearby sAZ synapses (arrowheads) clustered within 1.5 μm (dashed yellow ring). (B) Comparison of a normalized clustering index between CTB(+) dominant-eye mAZ and sAZ synapses at each time point in WT (left) and β2KO (right) mice. Black dots represent mean values from separate biological replicates and black lines connect measurements within each replicate (N=3 for each age and genotype). (C) Same as in (B) showing the normalized clustering index measurements for CTB(-) non-dominant-eye synapses. In B/C, paired T-tests were used to test the statistical significance between mAZ and sAZ clustering index measurements. Error bars represent means ± SEMs. *: p<0.05; **: p<0.01. (D) Cumulative distributions of VGluT2 volume for CTB(+) dominant-eye sAZ synapses near (<1.5 μm, black lines) or far from (>1.5 μm, red lines) like-eye mAZ synapses. Left panel: WT. Right panel: β2KO. (E) Same as in (D) showing the cumulative distributions of VGluT2 volume for CTB(-) non-dominant-eye sAZ synapses relative to like-eye mAZ synapses. The distributions in D/E show merged data across all developmental ages (P2/P4/P8; N=3 biological replicates at each time point). A nonparametric Kolmogorov-Smirnov test was used for statistical analysis. “n.s.” indicates no significant difference between near and far sAZ synapse distributions from the K-S test. Separate measurements of 5/95% confidence intervals show overlap of the two distributions in each case.

Revised version: (A) Representative Z-projection images of mAZ and sAZ inputs across ages and genotypes. Arrowheads point to individual Bassoon/Homer1 cluster pairs indicating release sites. (B) Representative CTB(+) dominant-eye (top panels) and CTB(-) non-dominant-eye (bottom panels) mAZ inputs in a WT P8 sample, showing synaptic (left panels), CTB (middle panels), and merged labels (right panels). Arrowheads point to individual Bassoon/Homer1 paired clusters. (C) Eye-specific mAZ (left) and sAZ (right) input density across development in WT (top panels) and β2KO mice (bottom panels). Black dots represent mean values from separate biological replicates and black lines connect eye-specific measurements within each replicate (N=3 for each age and genotype). Error bars represent group means ± SEMs. Statistical significance between eye-specific measurements was assessed using two-tailed paired T-tests at each age. *p<0.05, **p<0.01.

Added in the revised version: After assigning eye-of-origin to each retinogeniculate synapse, we found that the density of CTB(+) mAZ inputs was significantly higher than CTB(-) mAZ inputs at every age in WT mice (Fig. 2C, top left; Fig. S1, top). β2KO mice also developed eye-specific differences in mAZ input density (Fig. 2C, bottom left; Fig. S1, bottom). For WT sAZ synapses, the dominant-eye had a higher synapse density at P2 and P8. At P4, however, sAZ synapse density was equivalent between the eyes (Fig. 2C, top right).

Added in the revised version: This pattern resulted from non-dominant eye sAZ synapse addition, which equalized the global input density between the eyes at P4 as we reported previously (5/95% confidence interval,-0.014-0.011 synapses / μm3). In β2KO mice, the non-dominant eye formed fewer sAZ synapses, trending toward an eye-specific difference. (Fig. 2C, bottom right; p<0.05, power=0.62; see Tables S2/S3). Thus, β2KO mice with disrupted retinal activity maintain a higher mAZ input fraction in the dominant-eye projection, but sAZ synapse addition is reduced at the peak of competition.

Added in the revised version: Multi and single active zone inputs from the dominant eye show increased vesicle pool size and vesicle proximity to the active zone

Added in the revised version: We previously reported a dominant-eye bias in VGluT2 volume when considering all retinogeniculate inputs 26. Here, we assessed presynaptic maturation separately in sAZ and mAZ inputs by measuring their total VGluT2 volume. In WT mice, both mAZ (Fig. 3A, left) and sAZ (Fig. 3B, left) inputs showed significant eye-specific volume differences at each age. In the middle of eye-specific competition at P4 in WT mice, the median VGluT2 cluster volume in dominant-eye mAZ inputs was 3.72 fold larger than that of non-dominant-eye inputs (Fig. 3A, left). In contrast, β2KO mice showed a smaller 1.1 fold difference at the same age (Fig. 3A, right panel). For sAZ synapses at P4, the magnitudes of eye-specific differences in VGluT2 volume were smaller: 1.35-fold in WT (Fig. 3B, left) and 0.41-fold in β2KO mice (Fig. 3B, right). Thus, both mAZ and sAZ input size favors the dominant eye, with larger eye-specific differences seen in WT mice (see Table S3).

Added in the revised version: Dominant-eye inputs show larger vesicle pools that scale with active zone number.

Added in the revised version: (A and B) Violin plots showing the distribution of VGluT2 cluster volume for (A) mAZ and (B) sAZ inputs in WT (filled) and β2KO mice (striped) at each age. The width of each violin plot reflects the relative synapse proportions across the entire grouped data set at each age (N=3 biological replicates) and the maximum width was normalized across all groups. The black dots represent the median value of each biological replicate (N=3), and the black horizontal lines represent the median value of all inputs grouped across replicates. Black lines connect measurements of CTB(+) and CTB(-) populations from the same biological replicate. Statistical significance was determined using a linear mixed model ANOVA with a post-hoc Bonferroni correction. Black asterisks indicate significant eye-specific differences at each age. *p<0.05, **p<0.01. ***p<0.001. (C) Eye-specific VGluT2 signal volume for all inputs separated by number of AZs in WT (left panel) and β2KO mice (right panel) at P4. (D) Average VGluT2 volume per AZ for all inputs separated by number of AZs in WT (left panel) and β2KO mice (right panel) at P4. In panels (C) and (D), error bars indicate group means ± SEMs (N=3 biological replicates for each age and genotype). Black dots represent mean values from separate biological replicates and black lines connect eye-specific measurements within each replicate. Statistical significance between eye-specific measurements was assessed using two-tailed paired T-tests: *p<0.05; **p<0.01.

Added in the revised version: In addition to total vesicle pool volume, we quantified VGluT2 volume within a 70 nm shell around each AZ in mAZ and sAZ inputs, considering this a proxy for docked vesicles (Fig. S2A-C) 26. In WT mice at P4, dominant-eye inputs showed greater per-AZ vesicle volume than non-dominant inputs in both mAZ and sAZ categories (Fig. S2B, left). These eye-specific differences were absent in β2KO mice (Fig. S2B, right).

Added in the revised version: Moreover, at all ages and in both genotypes, vesicle volume within the 70 nm shell did not differ between mAZ and sAZ inputs (Fig. S2A-C). This shows that vesicle docking at release sites favors the dominant-eye as we previously reported 26, but is similar for like-eye type inputs regardless of AZ number.

Added in the revised version: Vesicle pool size scales with active zone number

Added in the revised version: Because mAZ inputs showed greater total VGluT2 volume than sAZ synapses, yet exhibited comparable vesicle docking, the disparity could reflect a scaling effect of vesicle pool size with increased AZ number. To evaluate how presynaptic vesicle pool volume scales with AZ number, we compared the Bassoon cluster number to VGluT2 volume for every retinogeniculate input. In both WT (Fig. S2D) and β2KO mice (Fig. S2E), mAZ inputs contained an average of 2-3 Bassoon clusters (separate AZs) in the first postnatal week. In WT mice, CTB(+) dominant-eye mAZ inputs contained more AZs than CTB(-) non-dominant-eye mAZ inputs at P4 and P8 (Fig. S2D). This maturation was delayed until P8 in β2KO mice (Fig. S2E). For CTB(+) dominant-eye inputs in both genotypes at P4, vesicle pool volume correlated positively with AZ number (Fig. 3C; Pearson correlation coefficients: WT [0.99] and β2KO [0.95]). A similar, but weaker correlation was observed for CTB(-) non-dominant-eye inputs (Pearson correlation coefficients: WT [0.97] and β2KO [0.90]. Dividing the total presynaptic VGluT2 volume by the AZ number revealed a consistent vesicle volume per AZ for both sAZ and mAZ inputs (Fig. 3D/S3, Table S3). Collectively, these results indicate that presynaptic vesicle pool volume scales with AZ number for each eye-specific input, while a dominant-eye bias persists throughout development.

Added in the revised version: Eye-specific synapse clustering before eye-opening

Added in the revised version: Eye-specific competition is thought to involve stabilization of co-active, neighboring inputs from the same eye and elimination of out of sync inputs from the opposite eye 42–44. Because axon refinement depends on the relative strength of neurotransmission between competing inputs 39, 45–51, mAZ inputs with multiple release sites could help stabilize nearby like-eye inputs 48, 52–54.

Added in the revised version: To quantify synaptic clustering patterns, we measured the distance from every eye-specific sAZ synapse to all other sAZ and mAZ inputs within each image field. sAZ synapses were often found nearby other inputs from the same eye (Fig. 4A). To define clustering, we searched volumetrically around each mAZ input and measured the fraction of sAZ synapses that were nearby at increasing distances (Fig. S4A/B). We then compared the observed percentages with datasets in which sAZ synapse positions were randomly shuffled within each neuropil volume. The largest differences occurred within a 1-2 μm search distance (Fig. S4A/B), and so we chose a 1.5 μm cutoff from the edge of each mAZ input to designate it “clustered” if at least one neighboring sAZ synapse lay within this distance and “isolated” if none did (Fig. 4B). No significant clustering was detected when sAZ and mAZ inputs originated from opposite eyes (Fig. S4C/D).

Added in the revised version: Eye-specific synapse clustering before eye-opening.

Added in the revised version: (A) Representative mAZ (left panels) and sAZ (right panels) inputs in a WT P8 sample with nearby sAZ synapses (arrowheads) clustered within 1.5 μm (dashed yellow ring). Arrows point to the centered mAZ or sAZ inputs. (B) Ratio of clustered and isolated mAZ and sAZ inputs for CTB(+) (upper panels) and CTB(-) (lower panels) inputs in WT and β2KO mice at P4. (C) Comparison of the clustered input ratio between mAZ and sAZ inputs across different ages, genotypes, and eyes of origin. (D) Comparison of the average number of nearby sAZ synapses for clustered mAZ and sAZ inputs across different ages, genotypes, and eyes of origin. In panels B-D, black dots represent mean values from separate biological replicates and black lines connect measurements within each replicate (N=3 for each age and genotype). Error bars represent group means ± SEMs. Two-tailed paired T-tests were used to test statistical significance between mAZ and sAZ inputs at each age. *p<0.05; **p<0.01.

Added in the revised version: At the peak of competition (P4) in WT mice, more than 65% of both mAZ and sAZ inputs were clustered for both eyes, while this proportion fell to ∼50% in β2KO mice (Fig. 4B). In the WT dominant-eye projection, the fractions of clustered mAZ and sAZ inputs were similar at each age (Fig. 4C, top left panel). For the non-dominant eye projection, however, clustered mAZ inputs outnumbered clustered sAZ inputs at P4 (Fig. 4C, bottom left panel), the age when this eye adds sAZ synapses (Fig. 2C). β2KO mice showed no difference between mAZ and sAZ clustering at any age (Fig. 4C, right panels). Additionally, in WT mice at P4 and P8, clustered mAZ inputs from both eyes had more neighboring sAZ synapses than did clustered sAZ synapses (Fig. 4D, left); this enrichment was absent in β2KO mice (Fig. 4D, right). Thus, while most retinogeniculate synapses lie within 1.5 μm of another like-eye input, WT mice show a relative increase in neighboring synapses near mAZ inputs during synaptic competition.

Added in the revised version: Clustered mAZ inputs in P4 WT mice were also closer together than isolated mAZ inputs. Clustered mAZ inputs formed by the dominant-eye were ∼32% closer to the nearest like-eye clustered mAZ input compared to isolated mAZ inputs (Fig. 5A, left panel). In the non-dominant eye projection, clustered mAZ inputs were ∼55% closer together compared to isolated mAZ inputs (Fig. 5B, left panel). Once segregation was complete at P8, distances between clustered and isolated mAZ inputs were more similar (Fig. S5). In β2KO mice, distances between isolated and clustered mAZ inputs and their nearest clustered mAZ neighbor did not differ for either eye at P4 or P8 (Fig. 5/S5, right panels). These patterns are consistent with sAZ synapse addition when mAZ inputs are in close proximity during competition.

Added in the revised version: Clustered mAZ inputs are closer than isolated inputs during competition.

Added in the revised version: (A-B) Distance between clustered and isolated mAZ inputs and the closest like-eye clustered mAZ input, shown for (A) CTB(+) and (B) CTB(-) projections at P4 in WT and β2KO mice. Boxes indicate the 25%-75% distribution of input measurements from N=3 biological replicates and whiskers extend to 1.5 times the interquartile range. Gray dots represent individual distance measurements for all mAZ inputs. Black and red dots represent mean values from separate biological replicates and black lines connect measurements within each replicate (N=3 for each age and genotype). Statistical significance was determined using a linear mixed model ANOVA with post-hoc Bonferroni correction. Black asterisks indicate significant differences. *p<0.05, ***p<0.001.

Added in the revised version: Finally, we tested whether sAZ vesicle pool volume varies with distance to mAZ inputs. We classified each sAZ as “near” (≤1.5 μm) or “far” (>1.5 μm) relative to the nearest like-eye mAZ input. Across ages and genotypes, vesicle pool volume was similar for near and far sAZ synapses in both eyes (Fig. S6). Thus, vesicle pool size in sAZ synapses appears independent of their proximity to mAZ inputs.

Removed in the revised version: In addition to sAZ clustering near mAZ synapses, sAZ synapses also showed a non-random relationship to other sAZ synapses across ages and genotypes (Figure S4A and S4B). Although axonal path lengths were not available in our dataset, our volume-based randomization allowed for direct comparison of clustering magnitudes between mAZ and sAZ synapses. We calculated a clustering index by dividing the “nearby” sAZ synapse ratio (% of total synapses) in the original data by that in the randomized data(Figure 3, B and C). This index revealed that for dominant-eye synapses, greater clustering occurred for mAZ synapses across all ages and genotypes (Figure 3B).

Removed in the revised version: However, for non-dominant eye synapses, the mAZ versus sAZ contrast was only found in WT mice at P4 (Figure 3C). At this stage, ipsilateral RGC axon synaptogenesis continues and ipsilateral sAZ synapses form near mAZ synapses despite being in the future contralateral eye-specific territory. However, after eye-specific competition at P8, non-dominant-eye CTB(-) sAZ synapses were eliminated nearby CTB(-) mAZ synapses from P4-P8, while the mAZ synapses were retained (the 5/95% confidence interval for the change in mAZ synapse density from P4 to P8 in β2KO mice was -0.0011∼0.0028 synapses / μm3) (Figure 1E).

Removed in the revised version: In β2KO mice, dominant-eye projections also showed increased clustering around mAZ synapses compared to sAZ synapses (Figure 3B, right panel), suggesting partial maintenance of synaptic clustering despite retinal wave defects. However, non-dominant eye mAZ synapses in β2KO mice did not show clustering at P4 (Figure 3C, right panel, 5/95% confidence interval: -0.033∼0.42), consistent with the overall reduction in sAZ synapse number compared with WT mice (Figure S1B). Together, these results show that defects in spontaneous retinal activity reduce synaptic clustering in an eye-specific manner.

Removed in the revised version: To validate the robustness of our comparison between original and randomized data across randomization seeds, we quantified the ratio of sAZ synapses nearby mAZ or sAZ center synapses in 10 randomized trials (Figure S4C, representative data from a WTP4 sample). The result showed that variation across 10 randomizations was minimal and did not impact the comparison.

Removed in the revised version: Considering that sAZ synapses show eye-specific differences in vesicle pool volume (Figure 2 A/B) and cluster nearby mAZ synapses (Figure 3 A-C), we hypothesized that mAZ synapses might influence the vesicle pool size of nearby sAZ synapses. To test this, we classified sAZ as “near” (within 1.5 μm of a like-eye mAZ synapse) or “far” (any distance >1.5 μm). Across all ages and genotypes, we found no significant differences in vesicle pool volume between near and far sAZ synapses for either dominant-eye (difference in VGluT2 volume: -0.08 ∼ 0.004 log10[μm3], 5/95% confidence interval) or non-dominant-eye synapses (difference in VGluT2 volume: -0.06 ∼ 0.105 log10[μm3], 5/95% confidence interval) (Figure 3 D/E, displays data merged for all time points). This result shows that sAZ vesicle pool size is independent of their proximity to mAZ synapses.

Removed in the revised version: Synaptic clustering varies with distance to neighboring eye-specific multi-active zone synapses

Removed in the revised version: RGC axon refinement is a dynamic process involving branch stabilization and elimination based on the relative activity patterns among neighboring inputs (41, 44). Axonal remodeling is regulated by synaptic transmission (37, 45, 46) and non-cell-autonomous stabilization and punishment signals (42, 43, 47). Although the precise mechanisms of axonal stabilization and punishment are not fully understood, it is likely that non-cell-autonomous signals operate at a local scale through direct cell-cell interactions or diffusible paracrine factors. Because CTB(-) non-dominant-eye sAZ synapses initially cluster around like-eye mAZ synapses and are later eliminated during eye-specific competition (Figure 3C), we hypothesized that these clustering effects may correlate with distance to the nearest neighboring eye-specific mAZ synapses.

Removed in the revised version: To explore this, we first examined whether the size of CTB(-) non-dominant-eye mAZ synapses varied based on their proximity to CTB(+) dominant-eye mAZ synapses (Figure 4A/B). For each CTB(-) mAZ synapse, we measured the vesicle pool volume (Figure 4A, P4 data as an example) and AZ number (Figure 4B) as a function of each synapse’s distance to the nearest CTB(+) mAZ synapse. We found no correlation between presynaptic properties and inter-eye mAZ synapse distances across all ages in both WT mice (Figure 4 A/B, left panels) and β2KO mice (Figure 4A/B, right panels) (P2/P8 data not shown). These results show that vesicle pool size and active zone number within non-dominant-eye mAZ synapses are not influenced by proximity to dominant-eye mAZ synapses.

Removed in the revised version: Synaptic clustering varies with distance to neighboring eye-specific multi-active zone synapses.

Removed in the revised version: (A) VGluT2 volume of CTB(-) non-dominant-eye mAZ synapses relative to their distance to the nearest CTB(+) dominant-eye mAZ synapse in a WT P4 sample (left panel) and a β2KO P4 sample (right panel). Each black dot represents one mAZ synapse. The red solid line represents the result of linear regression fitting. (B) Distributions of distances between CTB(-) non-dominant-eye mAZ synapses and their nearest CTB(+) dominant-eye mAZ synapse separated by the number of AZs within each CTB(-) mAZ synapse in WT P4 samples (left panel) and β2KO P4 samples (right panel). The median value is indicated by the horizontal line within the box, while the box boundaries represent quartile values. The whiskers represent the maximum and minimum values. A mixed model ANOVA was used to perform statistical tests. “n.s.” indicates no significant differences. The red solid line represents the result of linear regression fitting. (C) Cumulative distributions of distances between CTB(+) dominant-eye mAZ synapses and their nearest CTB(+) dominant-eye mAZ synapse (cartoon) in WT P4 samples (left panel) and β2KO P4 samples (right panel). Red lines indicate clustered mAZ synapses with nearby (<1.5 μm) sAZ synapses and black lines indicate isolated mAZ synapses with no nearby sAZ synapses. (D) Same presentation as in (C), showing distances between CTB(-) mAZ synapses. (E) Same presentation as in (C), showing distances between CTB(-) non-dominant-eye mAZ synapses and their nearest CTB(+) dominant-eye mAZ synapse. (F) Same presentation as in (C), showing distances between CTB(+) dominant-eye mAZ synapses and the nearest CTB(-) non-dominant-eye mAZ synapse. For C-F, nonparametric Kolmogorov-Smirnov tests were used for statistical comparisons. “***” indicates p<0.001, while “n.s.” indicates no significant difference from the K-S test results. Separate measurements of 5/95% confidence intervals show overlap of the two distributions in each “n.s.” instance.

Removed in the revised version: Next, we investigated whether synaptic clustering around mAZ synapses from one eye varied based on their distance to the nearest neighboring eye-specific mAZ synapses. For this analysis, we categorized eye-specific mAZ synapses into two groups: “clustered” mAZ synapses with nearby (< 1.5 μm) sAZ synapses from the same eye, and “isolated” mAZ synapses with no nearby sAZ synapses. For the dominant-eye projection, 62-73% of all mAZ synapses were clustered (P2-P8). The mAZ ratios for the non-dominant-eye projection were 46% (P2), 51% (P4), and 54% (P8) (Figure S5). After distinguishing these two types of mAZ synapses, we measured the distance between each type and its closest mAZ synapses from each eye-of-origin (Figure 4C-F).

Removed in the revised version: In WT mice at P4, CTB(+) dominant-eye clustered mAZ synapses were closer to other CTB(+) mAZ synapses compared to CTB(+) isolated mAZ synapses (Figure 4C, left panel). Similarly, CTB(-) non-dominant-eye clustered mAZ synapses were closer to the nearest CTB(-) mAZ synapse compared to CTB(-) isolated mAZ synapses (Figure 4D, left panel). These distance-dependent relationships were not observed when mAZ synapse positions of the target eye were randomized (Figure S6A/B, 5/95% confidence interval for clustered and isolated mAZ synapses: 4.40∼4.60 μm and 4.44∼4.60 μm from CTB(+) to CTB(+) synapses; 2.82∼3.08 μm and 2.90∼3.10 μm from CTB(+) to CTB(-) synapses; 4.30∼4.53 μm and 4.12∼4.40 μm from CTB(-) to CTB(-) synapses; 3.26∼3.49 μm and 3.26∼3.54 μm from CTB(-) to CTB(+) synapses). These findings indicate that synaptic clustering around a mAZ synapse is greater when it is near another mAZ synapse from the same eye. This distance-dependent clustering persisted at P8 for CTB(+) dominant-eye mAZ synapses (Figure S6C, left panel). However, at P8, CTB(-) non-dominant-eye mAZ synapses did not show a distance-dependent clustering effect (Figure S6D, right panel). This reflected a change in the ratio of clustered/isolated CTB(-) mAZ synapses from P4-P8 (Figure S5, 5/95% confidence interval in the difference was -50.5% ∼ -23.4%). Despite this shift, the density of CTB(-) mAZ synapses remained consistent (Figure 1E, 5/95% confidence interval in the difference was -0.05 ∼ 0.08 synapses / μm3).

Removed in the revised version: To look for evidence of competitive interactions between synapses from the two eyes, we measured the distances between clustered and isolated CTB(-) non-dominant-eye mAZ synapses and the closest CTB(+) dominant-eye mAZ synapse (Figure 4E). In WT mice at P4, CTB(-) non-dominant-eye isolated mAZ synapses were closer to the nearest CTB(+) dominant-eye mAZ synapse compared to CTB(-) clustered mAZ synapses (Figure 4E, left panel). This effect was not observed when CTB(+) mAZ synapse positions were randomized (Figure S6A/B).

Removed in the revised version: Lastly, we compared the distances between clustered and isolated CTB(+) dominant-eye mAZ synapses with respect to CTB(-) non-dominant-eye mAZ synapses and found no differences in WT mice (Figure 4F, left panel, the 5/95% confidence intervals: 3.92 ∼ 4.51 μm for clustered mAZ synapses and 3.98 ∼ 4.29 μm for isolated mAZ synapses). This suggests that dominant-eye synaptic clustering is not affected by proximity to non-dominant-eye mAZ synapses.

Removed in the revised version: For all measured inter-synaptic relationships, distance-dependent effects on synaptic clustering were only observed in WT mice and not in β2KO mice (Figure 4 C/D, right panels; the 5/95% confidence interval for distances were 6.75 ∼ 7.06 μm and 6.62 ∼ 6.99 μm for clustered and isolated CTB(+) mAZ synapses, and 4.51 ∼ 6.59 μm and 5.36 ∼ 6.43 μm for clustered and isolated CTB (-) mAZ synapses). Together, these results show that normal spontaneous retinal activity is necessary for non-random eye-specific synaptic clustering during synaptic competition.

Discussion

Added in the revised version: Spontaneous retinal activity guides eye-specific refinement by strengthening dominant eye inputs in the correct territory and pruning weaker inputs from the competing eye. Using volumetric super-resolution microscopy and eye-specific synaptic immunolabeling, we previously identified activity-dependent, eye-specific differences in presynaptic vesicle pool size and vesicle association with the active zone (AZ), both favoring the dominant eye during synaptic competition 26. Here, we reanalyzed this dataset to measure synaptic spatial arrangements during normal development and how these are affected in a mutant model with abnormal eye-specific segregation.

Added in the revised version: Early input maturation and addition of release sites

Added in the revised version: Before eye-opening, retinal ganglion cell (RGC) axons from both eyes formed terminals with either multiple active zones (mAZ inputs) or a single active zone (sAZ synapses). Electron microscopy of dAPEX2-labeled ipsilateral RGC inputs revealed individual boutons with multiple active zones. Prior electron microscopy studies in mice showed no clustering from neighboring RGC boutons at postnatal day 8 10, 11, suggesting the mAZ inputs seen in STORM images are single boutons containing multiple release sites. During competition, the dominant eye projection generated more of these mAZ inputs, each with more active zones and larger vesicle pools than the non-dominant eye. We observed a linear relationship between presynaptic vesicle signal volume and active zone number, suggesting that release site addition and vesicle pool expansion may be linked during input maturation. Functional recordings indicate that release site addition is the primary driver of retinogeniculate input strengthening after eye-opening 24. Our findings reveal that eye-specific differences in release site addition at individual terminals emerge at the earliest stages of binocular refinement.

Added in the revised version: Spatial interactions during competition

Added in the revised version: Ipsilateral RGC axons are delayed in entering the dLGN until after contralateral inputs have already innervated the nucleus 55. Despite their late arrival, these axons competed by forming numerous sAZ inputs, which equalized the overall input density between the eyes in the future contralateral eye domain at postnatal day 4 (P4). Most of these synapses formed nearby other like-eye neighbors, and were enriched near mAZ inputs, supporting the idea of release site addition at strong synapses. Neighboring sAZ synapses could ultimately mature into mAZ inputs as release sites are added and vesicle pools expand.

Added in the revised version: Distance analysis showed that clustered mAZ boutons were closer together than isolated mAZ boutons, further supporting release site addition near strong neighboring inputs. It remains unknown, however, whether nearby mAZ inputs originate from the same RGC. Nearby mAZ inputs within individual RGC arbors could promote nearby release site addition via intrinsic mechanisms, such as increased presynaptic calcium entry or surface delivery of synaptogenic molecules. Boutons with multiple release sites could also trigger non-cell-autonomous signaling mechanisms that stabilize neighboring inputs from the same eye 48, 53, 54. Release site addition increases extracellular glutamate concentration and promotes spillover onto adjacent synapses, which enhances the excitability of developing relay cells 25 and could contribute to long term synaptic plasticity evoked by high frequency spiking in retinal wave bursts 56–59. Consequently, release site addition and local bouton clustering may act in tandem to stabilize coactive inputs.

Removed in the revised version: Relay neurons of the dLGN function as critical integrators of locally clustered RGC inputs, which drive spike output to the primary visual cortex. In this study, we used volumetric super-resolution microscopy and eye-specific synaptic immunolabeling to show that 1) eye-specific synaptic clustering begins during retinogeniculate refinement, before eye-opening, 2)clustering occurs near RGC inputs with multiple active zones (mAZ synapses), 3)synaptic clustering depends on spontaneous retinal activity during the first postnatal week, and 4) the proximity of eye-specific mAZ synapses influences activity-dependent clustering, with nearby like-eye inputs showing increased clustering and opposite-eye inputs exhibiting reduced clustering as a function of distance to competing synapses. These results show that synaptic clustering begins at the earliest stages of eye-specific refinement and are consistent with previous studies implicating local signaling mechanisms in activity-dependent synaptic competition in the developing visual system.

Removed in the revised version: Previous anatomical studies have shown that retinogeniculate glomeruli mature through the progressive clustering of RGC boutons after eye-opening in mice (10, 13, 20). These clusters form through the convergence of inputs from multiple RGCs (9-14) and by the clustering of boutons within the arbors of individual RGCs (20). The latter depends on visual experience and is partly disrupted by sensory deprivation (e.g., late dark-rearing) (20). At maturity, retinogeniculate clusters are primarily eye-specific and contain very few, weak synapses from the non-dominant eye (15, 17). Our STORM images reveal that eye-specific synaptic clustering begins before eye-opening, during the period of eye-specific competition driven by stage II cholinergic retinal waves. This clustering is characterized by the emergence of a subset of synapses containing multiple AZs (mAZ synapses), distinct from isolated RGC terminals with a single active zone (sAZ synapses). Electron microscopy images of dAPEX2-labeled RGC inputs confirmed the presence of retinogeniculate synapses with multiple active zones. Prior electron microscopy studies in the mouse found limited evidence of convergent synaptic clustering from neighboring RGCs at postnatal day 8 (10, 13), suggesting that the mAZ synapses seen in STORM images are single retinogeniculate terminals. The lack of synaptic convergence in prior EM reconstructions at P8 implies that early clustering around mAZ synapses may result from local output clustering within individual RGC arbors.

Removed in the revised version: Our previous analysis of this STORM data set (25) demonstrated eye-specific differences in presynaptic vesicle pool size and vesicle association with the AZ as key factors in eye-specific competition. Here, we further identify several effects contributing to these differences: 1) Both sAZ and mAZ synapses exhibit eye-specific differences in VGluT2 volume. By measuring VGluT2 volume near presynaptic Bassoon clusters, we found that both synapse types show eye-specific differences in VGluT2 enrichment at the AZ (Figure 2A/B), while the extent of putative vesicle docking remains similar between synapse types for each eye (Figure S7, 5/95% confidence interval for differences in VGluT2 volume near the AZ was -0.63∼0.191 log10[μm3] for WT mice and -0.108∼0.129 log10[μm3] for β2KO mice). 2) Dominant-eye projections form more mAZ synapses than the non-dominant-eye projections at each stage of eye-specific refinement (Figure 1F). Moreover, mAZ synapses have ∼3 fold larger presynaptic vesicle pools than sAZ synapses at the peak of synaptic competition, when input densities are equivalent between the two eyes (Figure 2D). 3) On average, mAZ synapses from the dominant-eye contain more AZs per synapse compared to those from the non-dominant-eye throughout development (Figure 2C). These factors contribute to a greater eye-specific VGluT2 volume difference for mAZ synapses compared to sAZ synapses (Figure 2A/B). Since mAZ synapses are expected to have a higher release probability, they likely play an important role in driving plasticity mechanisms reliant on neurotransmission.

Removed in the revised version: Indeed, mAZ synapses act as local hubs for clustering like-eye-type sAZ synapses. These clustered synapses may be formed by the same RGC axon, or by adjacent RGC axons. Our current EM data, limited to two-dimensional images, leaves open the question of whether RGC synaptic convergence or local clustering within RGC arbors is the primary contributor. Future volumetric EM reconstructions will be needed to resolve this. Our spatial measurements from STORM images show that synaptic clustering is more likely when mAZ synapses from the same eye are closer together. This was observed for both dominant- and non-dominant-eye inputs during the peak of synaptic competition at P4 (Figure 4). Like-eye-type synaptic clustering is consistent with previous reports of non-cell-autonomous signaling mechanisms promoting eye-specific axonal stabilization (41, 42, 47). The formation of mAZ synapses may also promote clustering within individual RGC arbors, potentially through a combination of cell-intrinsic factors (e.g. elevated presynaptic calcium) and cell-cell signaling mechanisms (e.g. neurotransmission, surface delivery of transmembrane cargo) contributing to synaptic competition.

Earlier version: In addition to axon stabilization, competitive refinement also involves axonal retraction mediated by punishment signals. Genetic deletions of VGluT2 or RIM1 proteins in ipsilaterally-projecting RGCs decrease presynaptic vesicle release and prevent contralateral RGC axon retraction (punishment) from the ipsilateral eye-specific territory in the dLGN (37, 48). Consistent with a putative punishment mechanism, we found that non-dominant-eye mAZ synapses were less likely to be clustered if they were closer to a dominant-eye mAZ synapse. One downstream mediator of synaptic punishment is JAK2 kinase, which is phosphorylated in less active synapses (43). Similar to neurotransmission mutant phenotypes, disruption of JAK2 kinase signaling prevents axon retraction (punishment) during eye-specific competition (43).

Revised version: Competitive refinement also involves synapse elimination and axonal retraction through punishment signals. Genetic deletions of VGluT2 or RIM1 proteins in ipsilaterally-projecting RGCs decreased presynaptic vesicle release and prevented retraction of contralateral RGC axons from the ipsilateral territory 39, 45. One downstream mediator of synaptic punishment is JAK2 kinase, which is phosphorylated in less active synapses 52. Similar to neurotransmission mutant phenotypes, disruption of JAK2 signaling prevents axon retraction (punishment) during competition 52. The spatial analyses we developed here will enable future mapping of input-specific punishment signals during synaptic competition. This includes phospho-JAK2 as well as molecular tags for glial pruning of weak inputs during eye-specific segregation 60–62.

Added in the revised version: Requirement for spontaneous retinal activity

Added in the revised version: β2KO mice showed significant defects in synapse addition during the first postnatal week. mAZs still formed with similar overall ratios as in wild type controls (Fig. S1, Table S3) and eye-specific differences in vesicle pool size still emerged. However, β2KO mice failed to form many sAZ synapses at the height of competition, particularly in the late-arriving non-dominant (ipsilateral) eye projection. The failure to add synapses reduced synaptic clustering and more inputs formed in isolation in the mutants compared to controls.

Added in the revised version: While our findings support a role for spontaneous retinal activity in presynaptic release site addition and clustering, activity-dependent postsynaptic mechanisms also influence input refinement at later stages. Retinogeniculate synapses undergo postsynaptic strengthening and weakening through potentiation and depression mediated by AMPARs and NMDARs 56–59. After eye-specific segregation, spontaneous retinal activity is required for postsynaptic AMPAR insertion, synaptic strengthening, and elimination of weaker inputs 63. Continued maintenance of segregation depends on calcium influx into relay neurons via L-type calcium channels, further implicating postsynaptic signaling in late-stage refinement 64, 65. Input targeting may also be guided by molecular cues to form non-random, eye-specific connections with postsynaptic targets. Ipsilaterally-projecting RGCs have distinct gene expression profiles that specify axon guidance and may further support eye-specific synaptic targeting 66, 67. Spontaneous retinal activity may permit axons to read out molecular regulators of synaptogenesis, as previously shown for RGC axon retraction 68.

Added in the revised version: Release site addition as a general mechanism underlying synaptic competition

Added in the revised version: Early synaptic clustering during retinogeniculate development resembles other circuits where neural activity guides competitive synaptic and axonal remodeling. At the neuromuscular junction (NMJ), motor neuron terminals compete for control of a postsynaptic muscle fiber; a single motor axon input strengthens while competing axons are eliminated 69–71. Competition at the NMJ depends on inter-synaptic distance, with motor axons losing connections located near stronger competing synapses 69, 70. As winning terminals are enlarged, presynaptic release site addition maintains a consistent density of Bassoon clusters (∼2-3/μm2) 72. Neighboring inputs with weaker synaptic transmission are selectively eliminated 73, 74. Competing motor axons differ in their release probabilities early in development 75, suggesting that presynaptic efficacy triggers local “stabilization” signals in winners and “punishment” signals in losers 76.

Added in the revised version: Presynaptic agrin release stabilizes postsynaptic receptor clusters by counteracting the destabilizing, dispersal effects of acetylcholine release, emphasizing the importance of early presynaptic transmission in postsynaptic stabilization 77, 78.

Added in the revised version: Similarly, in the developing cerebellum, Purkinje cells initially receive inputs from multiple climbing fibers (CFs) during the first postnatal week. Subsequently, one winning CF emerges and consolidates its synaptic inputs, while losing CFs are pruned 79, 80.

Added in the revised version: Here again, the winning input strengthens by adding presynaptic release sites, which increases multivesicular release and elevates glutamate concentration in the synaptic cleft 81–83. Postsynaptic ultrastructural and molecular changes occur several days later as dominant CF inputs potentiate and losing CF inputs depress 82, 84, 85 through spike-timing dependent remodeling 86, 87. Across these models, the earliest distinguishing feature between competing inputs is relative presynaptic transmission strength. Thus, release-site addition may be a conserved mechanism that biases synaptic refinement outcomes across developing neural circuits.

Removed in the revised version: In the adult dLGN, bouton clustering facilitates the convergence of RGC inputs representing similar visual features (e.g. direction of motion) and the integration of parallel visual channels carried by unique RGC mosaics (12, 15, 18). During development, the co-activation of neighboring RGCs by retinal waves may support the topographic refinement of clustered inputs to relay neurons. Our previous work demonstrated that abnormal retinal activity results in a significant reduction in eye-specific synapse numbers in β2KO mice (25). Additionally, β2KO mice show delayed eye-specific vesicle recruitment to the active zone by P8, a process present in WT mice as early as P2 (24). Similarly, our current analysis revealed increased vesicle pool volume and AZ number in β2KO mAZ synapses at later stages (P4 and P8) compared to WT mice (Figure 2). β2KO mice also lacked distance-dependent synaptic clustering effects until P8 (Figure S6E). Together, these results show how defects in retinal waves impact retinofugal synapse development, ultimately leading to abnormal receptive field properties (31, 49).

Removed in the revised version: Although our findings support the role of non-cell autonomous signaling in synaptic competition, alternative mechanisms may also contribute to eye-specific clustering. Input targeting could be molecularly specified to produce non-random eye-specific connections with postsynaptic targets. Ipsilaterally-projecting RGCs have a distinct gene expression profile (50), which may contribute to eye-specific targeting at the synaptic level.

Removed in the revised version: Our observation of synaptic clustering in the developing retinogeniculate system bears similarities to development in other circuits, such as the neuromuscular junction (NMJ). Motor neuron terminals undergo competitive refinement, where connections from a single motor axon are strengthened, while competing axons are eliminated (51-54). Synaptic competition at the NMJ is also dependent on inter-synaptic distance, with motor axons losing connections that are in closer proximity to competing synapses (51, 52). This activity-dependent process is biased toward the elimination of silenced inputs (55, 56). Interestingly, competing motor axon outputs show differences in presynaptic release probability at early stages of development, suggesting that biases in presynaptic release may underlie competition (57), possibly through the induction of local stabilization and punishment signals (54).

Removed in the revised version: Although the molecular mechanisms of eye-specific retinogeniculate clustering are still unknown, our STORM imaging results provide anatomical support for local, non-cell-autonomous interactions involved in synaptic refinement. Neighboring synapses may engage in direct signaling between presynaptic RGC axon terminals or initiate postsynaptic responses that lead to retrograde signaling to stabilize or eliminate synapses based on input timing. JAK2/STAT signaling is one known mediator of synapse elimination (43), guiding future investigations into the upstream signals that regulate competition. It will also be of interest to further investigate the synapse-specific induction of synaptic tags that mediate glial-dependent pruning of weak inputs during eye-specific segregation (58-60). Future studies using dense connectomic reconstruction and cell-type-specific markers will generate valuable anatomical snapshots of developing retinogeniculate circuits. An important goal will be to complement these anatomical resources with live nanoscopic imaging of rapid structural dynamics driving circuit refinement.

Materials and methods

Earlier version: The raw imaging data in this paper were previously reported (25). Materials and methods below are adapted from this work. All Matlab and Python code used in the work is available on GitHub (https://github.com/SpeerLab). Raw STORM images of the full data are available on the open-access Brain Imaging Library (61). These images can be accessed here: https://api.brainimagelibrary.org/web/view?bildid=ace-dud-lid (WTP2A data for example). All 18 datasets are currently searchable on the BIL by keyword “dLGN” or PI last name “Speer” and a DOI for the full data set is pending.

Revised version: The raw imaging data in this paper were previously reported 26. Materials and methods below are adapted from this work. All Matlab and Python code used in the work is available on GitHub (https://github.com/SpeerLab). Raw STORM images of the full data are available on the open-access Brain Imaging Library 88. These images can be accessed here: https://doi.brainimagelibrary.org/doi/10.35077/g.1164.

Animals

Earlier version: Wild-type C57BL/6J mice (Stock Number 000664) and APEX2 mice (Stock Number 032764) used in this study were purchased from the Jackson Laboratory. ET33Cre mice and β2KO were generously gifted by Drs. Eric M. Ullian (University of California, San Francisco) and Michael C. Crair (Yale School of Medicine), respectively. All experimental procedures were performed in accordance with an animal study protocol approved by the Institutional Animal Care and Use Committee (IACUC) at the University of Maryland. Neonatal male and female mice were used interchangeably for all experiments. Tissue from biological replicates (N=3 animals) was collected for each age (P2/P4/P8) from each genotype (WT and β2KO) (18 animals total). Primers used for genotyping β2KO mice are: forward: CAGGCGTTATCCACAAAGACAGA; reverse: TTGAGGGGAGCAGAACAGAATC; mutant reverse: ACTTGGGTTTGGGCGTGTTGAG (65, 66). Primers used for genotyping ET33Cre mice are: Forward: GGTCCTTGGCAGATGGGCAT; reverse: CGGCAAACGGACAGAAGCATT.

Revised version: Wild-type C57BL/6J mice (Stock Number 000664) and APEX2 mice (Stock Number 032764) used in this study were purchased from the Jackson Laboratory. ET33Cre mice and β2KO were generously gifted by Drs. Eric M. Ullian (University of California, San Francisco) and Michael C. Crair (Yale School of Medicine), respectively. All experimental procedures were performed in accordance with an animal study protocol approved by the Institutional Animal Care and Use Committee (IACUC) at the University of Maryland. Neonatal male and female mice were used interchangeably for all experiments. Tissue from biological replicates (N=3 animals) was collected for each age (P2/P4/P8) from each genotype (WT and β2KO) (18 animals total). Primers used for genotyping β2KO mice are: forward: CAGGCGTTATCCACAAAGACAGA; reverse:

Added in the revised version: TTGAGGGGAGCAGAACAGAATC; mutant reverse: ACTTGGGTTTGGGCGTGTTGAG.

Added in the revised version: Primers used for genotyping ET33Cre mice are: Forward: GGTCCTTGGCAGATGGGCAT; reverse: CGGCAAACGGACAGAAGCATT.

Electron microscopy tissue preparation

Earlier version: Animals were deeply anesthetized with ketamine/xylazine and transcardially perfused with 5-10 mls of 37°C 0.9% sterile saline (pH 7.2) followed by 20-30 mls of 37°C 4% EM-Grade paraformaldehyde (PFA, Electron Microscopy Sciences), 2% EM-Grade glutaraldehyde (GA, Electron Microscopy Sciences), 4mM calcium chloride (CaCl2, Sigma-Aldrich) in 0.2M cacodylate buffer (pH 7.4). Brains were postfixed in the same perfusion fixative solution overnight at 4°C. Brains were sectioned in 0.2M cacodylate buffer (pH 7.4) at 100 μm using a vibratome. A circular tissue punch (∼500 μm diameter) containing the dLGN was microdissected from each section using a blunt-end needle. dLGN sections were washed in 0.2M cacodylate buffer (pH 7.4) at 4°C on a rotator (4x 20 minutes). Sections were incubated in 20mM glycine (Sigma-Aldrich) in 0.2M cacodylate buffer (pH 7.4) at 4°C on a rotator for 30 minutes, followed by 5 x 20 minutes washes in 0.2M cacodylate buffer (pH 7.4) at 4°C on a rotator. dLGN sections were immersed in 0.5mg/mL DAB solution, covered with foil, for 30 minutes at 4°C on a rotator. 10 μL of 0.03% hydrogen peroxide (Sigma-Aldrich) was mixed into the DAB solution (1 mL) and incubated for 10 minutes at 4°C on a rotator (light protected). The reaction was quenched with 0.2M cacodylate buffer (pH 7.4) washes (3 x 1-minute washes followed by 2 x 20 minutes washes).

Revised version: Animals were deeply anesthetized with ketamine/xylazine and transcardially perfused with 5-10 mls of 37°C 0.9% sterile saline (pH 7.2) followed by 20-30 mls of 37°C 4% EM-Grade paraformaldehyde (PFA, Electron Microscopy Sciences), 2% EM-Grade glutaraldehyde (GA, Electron Microscopy Sciences), 4mM calcium chloride (CaCl2, Sigma-Aldrich) in 0.2M cacodylate buffer (pH 7.4). Brains were postfixed in the same perfusion fixative solution overnight at 4°C. Brains were vibratome sectioned in 0.2M cacodylate buffer (pH 7.4) at 100 μm. A circular tissue punch (∼500 μm diameter) containing the dLGN was microdissected from each section using a blunt-end needle. dLGN sections were washed in 0.2M cacodylate buffer (pH 7.4) at 4°C on a rotator (4x 20 minutes each). Sections were incubated in 20mM glycine (Sigma-Aldrich) in 0.2M cacodylate buffer (pH 7.4) at 4°C on a rotator for 30 minutes, followed by 5 x 20 minutes washes in 0.2M cacodylate buffer (pH 7.4) at 4°C on a rotator. dLGN sections were immersed in 0.5mg/mL DAB solution, covered with foil, for 30 minutes at 4°C on a rotator. 10 μL of 0.03% hydrogen peroxide (Sigma-Aldrich) was mixed into the DAB solution (1 mL) and incubated for 10 minutes at 4°C on a rotator (light protected). The reaction was quenched with 0.2M cacodylate buffer (pH 7.4) washes (3 x 1-minute washes followed by 2 x 20 minutes washes).

Tissue preparation for scanning electron microscopy

dLGN sections were fixed in 2% osmium tetroxide (Electron Microscopy Sciences) in 0.15M cacodylate buffer (pH 7.4) for 45 minutes at room temperature. Sections were reduced with 2.5% potassium ferricyanide (Electron Microscopy Sciences) in 0.15M cacodylate buffer (pH 7.4) for 45 minutes at room temperature in the dark, followed by 2 x 10-minute washes in double distilled water. Sections were incubated in 1% aqueous thiocarbohydrazide (Electron Microscopy Sciences) for 20 minutes, followed by 2 x 10-minute washes in double distilled water. Samples were fixed in 1% aqueous osmium tetroxide for 45 minutes at room temperature, followed by 2 x 10-minute washes in double distilled water. Sections were postfixed in 1% uranyl acetate (Electron Microscopy Sciences) in 25% ethanol in the dark for 20 minutes at room temperature and washed and stored in double distilled water overnight. The next day, samples were stained with aqueous lead aspartate at 60°C for 30 minutes and washed with double distilled water 2 x 10 minutes. Tissues were dehydrated in a graded dilution series of 100% ethanol (35%; 50%; 70%; 95%; 100%; 100%; 100% EtOH) for 10 minutes each at room temperature. Samples were immersed in propylene oxide (Electron Microscopy Sciences) 3 x 10 minutes at room temperature, followed by a series of epon resin/propylene oxide (812 Epon Resin, Electron Microscopy Sciences) exchanges with increasing resin concentrations (50% resin/50% propylene oxide; 65% resin/35% propylene oxide; 75% resin/25% propylene oxide; 100% resin; 100% resin) for 90 minutes each. Tissues were transferred to BEEM capsules (Electron Microscopy Sciences) that were filled with 100% resin and polymerized for at least 48 hours at 60°C.

Transmission electron microscopy image acquisition

Plasticized sections were cut at 70 nm with a Histo Jumbo diamond knife (DiATOME) using a Leica UC7 ultramicrotome. Sections were decompressed using chloroform vapor and collected onto 3.05mm 200 mesh Gilder thin bar hexagonal mesh copper grids (T200H-Cu, Electron Microscopy Sciences). Sections were imaged unstained on a Hitachi HT7700 transmission electron microscope (HT7700, Hitachi High-Tech America, Inc.) at 80kV.

Eye injections

Intraocular eye injections were performed one day before tissue collection. Briefly, mice were anesthetized by inhalant isoflurane and sterile surgical spring scissors were used to gently part the eyelid to expose the corneoscleral junction. A small hole was made in the eye using a sterile 34-gauge needle and ∼0.5 μL of cholera toxin subunit B conjugated with Alexa Fluor 488 (CTB-488, ThermoFisher Scientific, Catalogue Number: C34775) diluted in 0.9% sterile saline was intravitreally pressure-injected into the right eye using a pulled-glass micropipette coupled to a Picospritzer (Parker Hannifin).

dLGN tissue preparation for STORM imaging

Added in the revised version: Animals were deeply anesthetized with ketamine/xylazine and transcardially perfused with 5-10 mls of 37°C 0.9% sterile saline followed by 10 mls of room temperature 4% EM-Grade paraformaldehyde (PFA, Electron Microscopy Sciences) in 0.9% saline.

Earlier version: Animals were deeply anesthetized with ketamine/xylazine and transcardially perfused with 5-10 mls of 37°C 0.9% sterile saline followed by 10 mls of room temperature 4% EM-Grade paraformaldehyde (PFA, Electron Microscopy Sciences) in 0.9% saline. Brains were embedded in 2.5% agarose and sectioned in the coronal plane at 100 μm using a vibratome. From the full anterior-posterior series of dLGN sections (∼6-8 sections) we selected the central two sections for staining in all biological replicates. These sections were morphologically consistent with Figures 134-136 (5.07-5.31 mm) of the postnatal day 6 mouse brain from Paxinos, et al., “Atlas of the developing mouse brain” Academic Press, 2020 (62). Selected sections were postfixed in 4% PFA for 30 minutes at room temperature and then washed for 30-40 minutes in 1X PBS. The dLGN was identified by the presence of CTB-488 signals using a fluorescence dissecting microscope. A circular tissue punch (∼500 μm diameter) containing the dLGN was microdissected from each section using a blunt-end needle. A small microknife cut was made at the dorsal edge of the dLGN which, together with the CTB-488 signal, enabled us to identify the dLGN orientation during image acquisition.

Revised version: Brains were embedded in 2.5% agarose and vibratome sectioned in the coronal plane at 100 μm. From the full anterior-posterior series of dLGN sections (∼6-8 sections) we selected the central two sections for staining in all biological replicates. These sections were morphologically consistent with Figs. 134-136 (5.07-5.31 mm) of the postnatal day 6 mouse brain from Paxinos’s “Atlas of the developing mouse brain” Academic Press, 2020 89. Selected sections were postfixed in 4% PFA for 30 minutes at room temperature and washed for 30-40 minutes in 1X PBS. The dLGN was identified by the presence of CTB-488 signals using a fluorescence dissecting microscope. A circular tissue punch (∼500 μm diameter) containing the dLGN was microdissected from each section using a blunt-end needle. A small microknife cut was made at the dorsal edge of the dLGN which, together with the CTB-488 signal, enabled us to identify the dLGN orientation during image acquisition.

Immunohistochemistry

We used a serial-section single-molecule localization imaging approach to prepare samples and collect super-resolution fluorescence imaging volumes as previously described 26. dLGN tissue punches were blocked in 10% normal donkey serum (Jackson ImmunoResearch, Catalogue Number: 017-000-121) with 0.3% Triton X-100 (Sigma-Aldrich Inc.) and 0.02% sodium azide (Sigma-Aldrich Inc.) diluted in 1X PBS for 2-3 hours at room temperature and then incubated in primary antibodies for ∼72 hours at 4°C. Primary antibodies used were Rabbit anti-Homer1 (Synaptic Systems, Catalogue Number: 160003, 1:100) to label postsynaptic densities (PSDs), mouse anti-Bassoon (Abcam, Catalogue Number AB82958, 1:100) to label presynaptic active zones (AZs), and guinea pig anti-VGluT2 (Millipore, Catalogue Number AB251-I, 1:100) to label presynaptic vesicles. Following primary antibody incubation, tissues were washed in 1X PBS for 6 x 20 minutes at room temperature and incubated in secondary antibody solution overnight for ∼36 hours at 4°C. The secondary antibodies used were donkey anti-rabbit IgG (Jackson ImmunoResearch, Catalogue Number 711-005-152, 1:100) conjugated with Dy749P1 (Dyomics, Catalogue Number 749P1-01) and Alexa Fluor 405 (ThermoFisher, Catalogue Number: A30000), donkey anti-mouse IgG (Jackson ImmunoResearch, Catalogue Number 715-005-150, 1:100) conjugated with Alexa Fluor 647 (ThermoFisher, Catalogue Number: A20006) and Alexa Fluor 405, and donkey anti-guinea pig IgG (Jackson ImmunoResearch, Catalogue Number 706-005-148, 1:100) conjugated with Cy3b (Cytiva, Catalogue Number: PA63101). Tissues were washed 6 x 20 minutes in 1X PBS at room temperature after secondary antibody incubation.

Postfixation, dehydration, and embedding in epoxy resin

Tissue embedding was performed as previously described 26. Tissues were postfixed with 3% PFA + 0.1% GA (Electron Microscopy Sciences) in PBS for 2 hours at room temperature and then washed in 1X PBS for 20 minutes. To plasticize the tissues for ultrasectioning, the tissues were first dehydrated in a graded dilution series of 100% ethanol (50%/70%/90%/100%/100% EtOH) for 15 minutes each at room temperature and then immersed in a series of epoxy resin/100% EtOH exchanges (Electron Microscopy Sciences) with increasing resin concentration (25% resin/75% ethanol; 50% resin/50% ethanol; 75% resin/25% ethanol; 100% resin; 100% resin) for 2 hours each. Tissues were transferred to BEEM capsules (Electron Microscopy Sciences) that were filled with 100% resin and polymerized for 16 hours at 70°C.

Ultrasectioning

Plasticized tissue sections were cut using a Leica UC7 ultramicrotome at 70 nm using a Histo Jumbo diamond knife (DiATOME). Chloroform vapor was used to reduce compression after cutting. For each sample, ∼100 sections were collected on a coverslip coated with 0.5% gelatin and 0.05% chromium potassium (Sigma-Aldrich Inc.), dried at 60 degrees for 25 minutes, and protected from light prior to imaging.

Imaging chamber preparation

Coverslips were chemically etched in 10% sodium ethoxide for 5 minutes at room temperature to remove the epoxy resin and expose the dyes to the imaging buffer for optimal photoswitching. Coverslips were then rinsed with ethanol and dH2O. To create fiducial beads for flat-field and chromatic corrections, we mixed 715/755nm and 540/560nm, carboxylate-modified microspheres (Invitrogen, Catalogue Numbers F8799 and F8809, 1:8 ratio respectively) to create a high-density fiducial marker and then further diluted the mixture at 1:750 with Dulbecco’s PBS to create a low-density bead solution. Both high-and low-density bead solutions were spotted on the coverslip (∼0.7 µL each) for flat-field and chromatic aberration correction respectively. Excess beads were rinsed away with dH2O for 1-2 minutes. The coverslip was attached to a glass slide with double-sided tape to form an imaging chamber. The chamber was filled with STORM imaging buffer (10% glucose, 17.5µM glucose oxidase, 708nM catalase, 10mM MEA, 10mM NaCl, and 200mM Tris) and sealed with epoxy.

Imaging setup

Imaging was performed using a custom single-molecule super-resolution imaging system. The microscope contained low (4x/10x air) and high (60x 1.4NA oil immersion) magnitude objectives mounted on a commercial frame (Nikon Ti-U) with back optics arranged for oblique incident angle illumination. We used continuous-wave lasers at 488nm (Coherent), 561nm (MPB), 647nm (MPB), and 750nm (MPB) to excite Alexa 488, Cy3B, Alexa 647, and Dy749P1 dyes respectively. A 405 nm cube laser (Coherent) was used to reactivate Dy749P1 and Alexa647 dye photoswitching. The microscope was fitted with a custom pentaband/pentanotch dichroic filter set and a motorized emission filter wheel. The microscope also contained an IR laser-based focus lock system to maintain optimal focus during automatic image acquisition. Images were collected on 640*640-pixel region of an sCMOS camera (ORCA-Flash4.0 V3, Hamamatsu Photonics) with a pixel size of ∼155 nm.

Automated image acquisition

Fiducials and tissue sections on the coverslip were imaged using the low magnification objective (4X) to create a mosaic overview of the specimen. Beads/sections were then imaged at high magnification (60X) to select regions of interest (ROIs) in the Cy3B and Alexa 488 channels. Before final image acquisition, laser intensities and the incident angle were adjusted to optimize photoswitching for STORM imaging and utilize the full dynamic range of the camera for conventional imaging.

Low-density bead images were taken in 16 partially overlapping ROIs. 715/755nm beads were excited using 750 nm light and images were collected through Dy749P1 and Alexa 647 emission filters. 540/560nm beads were excited using a 488 nm laser and images were collected through Alexa 647, Cy3B, and Alexa 488 emission filters. These fiducial images were later used to generate a non-linear warping transform to correct chromatic aberration. Next, ROIs within each tissue section were imaged at conventional (diffraction-limited) resolution in all four-color channels sequentially.

Following conventional image acquisition, a partially overlapping series of 9 images were collected in the high-density bead field for all 4 channels (Dy749P1, Alexa 647, Cy3B, and Alexa 488). These images were later used to perform a flat-field image correction of non-uniform laser illumination across the ROIs. Another round of bead images was taken as described above in a different ROI of the low-density bead field. These images were later used to confirm the stability of chromatic offsets during imaging. All ROIs within physical sections were then imaged by STORM for Dy749P1 and Alexa 647 channels. Images were acquired using a custom progression of increasing 405nm laser intensity to control single-molecule switching. 8000 frames of Dy749P1 channel images were collected (60 Hz imaging) followed by 12000 frames of Alexa 647 channel images (100 Hz). In a second imaging pass, the same ROIs were imaged for Cy3B and Alexa 488 channels, each for 8000 frames (60 Hz).

We imaged the ipsilateral and contralateral ROIs separately in each physical section of the dLGN. For consistency of ROI selection across biological replicates at each age, we identified the dorsal-ventral (DV) axis of the dLGN and selected ROIs within the center (core region) at 2/5 (ipsilateral ROI) and 4/5 (contralateral ROI) of the full DV length.

Image processing

Earlier version: Single-molecule localization was performed using a previously described DAOSTORM algorithm (63) modified for use with sCMOS cameras (64). Molecule lists were rendered as 8-bit images with 15.5 nm pixel size where each molecule is plotted as an intensity distribution with an area reflecting its localization precision. Low-density fiducial images were used for chromatic aberration correction. We localized 715/755 beads in Dy749P1 and Alexa 647 channels, and 540/560 beads in Alexa 647, Cy3B, and Alexa 488 channels. A third-order polynomial transform map was generated by matching the positions of each bead in all channels to the Alexa 647 channel. The average residual error of bead matching was <15 nm for all channels. The transform maps were applied to both 4-color conventional and STORM images. Conventional images were upscaled (by 10X) to match the STORM image size. The method to align serial sections was previously described (25). STORM images were first aligned to their corresponding conventional images by image correlation. To generate an aligned 3D image stack from serial sections, we normalized the intensity of all Alexa 488 images and used these normalized images to generate both rigid and elastic transformation matrices for all four-color channels of both STORM and conventional data. The final image stack was then rotated and cropped to exclude incompletely imaged edge areas. To further confirm that the processed region corresponds to the contralateral dLGN, conventional CTB(+) signals of the labeled contralateral projection were thresholded to create a polygonal mask (a convex hull analysis linking the outermost CTB signals in the image volume). The mask was then applied to STORM images to exclude peripheral areas where CTB signals were absent or faint.

Revised version: Single-molecule localization was performed using a previously described DAOSTORM algorithm modified for use with sCMOS cameras 90, 91. Molecule lists were rendered as 8-bit images with 15.5 nm pixel size where each molecule is plotted as an intensity distribution with an area reflecting its localization precision. Low-density fiducial images were used for chromatic aberration correction. We localized 715/755 beads in Dy749P1 and Alexa 647 channels, and 540/560 beads in Alexa 647, Cy3B, and Alexa 488 channels. A third-order polynomial transform map was generated by matching the positions of each bead in all channels to the Alexa 647 channel. The average residual error of bead matching was <15 nm for all channels. The transform maps were applied to both 4-color conventional and STORM images. Conventional images were upscaled (by 10X) to match the STORM image size. The method to align serial sections was previously described 26. STORM images were first aligned to their corresponding conventional images by image correlation. To generate an aligned 3D image stack from serial sections, we normalized the intensity of all Alexa 488 images and used these normalized images to generate both rigid and elastic transformation matrices for all four-color channels of both STORM and conventional data. The final image stack was then rotated and cropped to exclude incompletely imaged edge areas. To further confirm that the processed region corresponds to the contralateral dLGN, conventional CTB(+) signals of the labeled contralateral projection were thresholded to create a polygonal mask (a convex hull analysis linking the outermost CTB signals in the image volume).

Added in the revised version: The mask was then applied to STORM images to exclude peripheral areas where CTB signals were absent or faint.

Cell body filter

Earlier version: The aligned STORM images had non-specific labeling of cell bodies in Dy749P1 and Alexa 647 channels corresponding to Homer1 and Bassoon immunolabels. To limit synaptic cluster identification to the neuropil region we identified cell bodies based on their Dy749P1 signal and excluded these regions from further image processing. STORM images were convolved with a Gaussian function (σ=140 nm) and then binarized using the lower threshold of a two-level Otsu threshold method. We located connected components in the thresholded images and generated a mask based on components larger than e11 voxels. Because cell body clusters were orders of magnitude larger than synaptic clusters, the cell body filter algorithm was robust to a range of size thresholds. The mask was applied to images of all channels to exclude cell body areas.

Revised version: The aligned STORM images had non-specific labeling of cell bodies in Dy749P1 and Alexa 647 channels corresponding to Homer1 and Bassoon immunolabels. To limit synaptic cluster identification to the neuropil region we identified cell bodies based on their Dy749P1 signal and excluded these regions from further image processing.

Added in the revised version: STORM images were convolved with a Gaussian function (σ=140 nm) and then binarized using the lower threshold of a two-level Otsu threshold method. We located connected components in the thresholded images and generated a mask based on components larger than e11 voxels. Because cell body clusters were orders of magnitude larger than synaptic clusters, the cell body filter algorithm was robust to a range of size thresholds. The mask was applied to images of all channels to exclude cell body areas.

Eye-specific synapse identification and quantification

Added in the revised version: To correct for minor variance in image intensity across physical sections, we normalized the pixel intensity histogram of each section to the average histogram of all sections.

Earlier version: To correct for minor variance in image intensity across physical sections, we normalized the pixel intensity histogram of each section to the average histogram of all sections. Image histograms were rescaled to make full use of the 8-bit range. Using a two-level Otsu threshold method, the conventional images were thresholded into three classes: a low-intensity background, low-intensity signals above the background representing non-synaptic labeling, and high-intensity signals representing synaptic structures. The conventional images were binarized by the lower two-level Otsu threshold, generating a mask for STORM images to filter out background signals. STORM images were convolved with a Gaussian function (σ= 77.5 nm) and thresholded using the higher two-level Otsu threshold. Following thresholding, connected components were identified in three dimensions using MATLAB ‘conncomp’ function. A watershedding approach was applied to split large clusters that were improperly connected. Clusters were kept for further analysis only if they contained aligned image information across two or more physical sections. We also removed all edge synapses from our analysis by excluding synapses that did not have blank image data on all adjacent sides.

Revised version: Image histograms were rescaled to make full use of the 8-bit range. Using a two-level Otsu threshold method, the conventional images were thresholded into three classes: a low-intensity background, low-intensity signals above the background representing non-synaptic labeling, and high-intensity signals representing synaptic structures. The conventional images were binarized by the lower two-level Otsu threshold, generating a mask for STORM images to filter out background signals. STORM images were convolved with a Gaussian function (σ= 77.5 nm) and thresholded using the higher two-level Otsu threshold. Following thresholding, connected components were identified in three dimensions using MATLAB ‘conncomp’ function. A watershedding approach was applied to split large clusters that were improperly connected. Clusters were kept for further analysis only if they contained aligned image information across two or more physical sections. We also removed all edge synapses from our analysis by excluding synapses that did not have blank image data on all adjacent sides.

Earlier version: To distinguish non-specific immunolabeling from true synaptic signals, we quantified two parameters for each cluster: cluster volume and cluster signal density calculated by the ratio of within-cluster pixels with positive signal intensity in the raw STORM images. Two separate populations were identified in 2D histograms plotted from these two parameters. We manually selected the population with higher volumes and signal densities representing synaptic structures. To test the robustness of the manual selection, we performed multiple repeated measurements of the same data and discovered a between-measurement variance of <1% (data not shown).

Revised version: To distinguish non-specific immunolabeling from true synaptic signals, we quantified two parameters for each cluster: cluster volume and cluster signal density calculated by the ratio of within-cluster pixels with positive signal intensity in the raw STORM images.

Added in the revised version: Two separate populations were identified in 2D histograms plotted from these two parameters. We manually selected the population with higher volumes and signal densities representing synaptic structures. To test the robustness of the manual selection, we performed multiple repeated measurements of the same data and discovered a between-measurement variance of <1% (data not shown).

Earlier version: To identify paired pre- and postsynaptic clusters, we first measured the centroid-centroid distance of each cluster in the Dy749P1 (Homer1) and Alexa 647 (Bassoon) channels to the closest cluster in the other channel. We next quantified the signal intensity of each opposing synaptic channel within a 140 nm shell surrounding each cluster. A 2D histogram was plotted based on the measured centroid-centroid distances and opposing channel signal densities of each cluster. Paired clusters with closely positioned centroids and high intensities of apposed channel signal were identified using the OPTICS algorithm. In total we identified 49,414 synapses from WT samples (3 samples each at P2/P4/P8, 9 total samples) and 33,478 synapses in β2-/- mutants (3 samples each at P2/P4/P8, 9 total samples). Retinogeniculate synapses were identified by pairing Bassoon (Alexa 647) clusters with VGluT2 (Cy3B) clusters using the same method as pre/post-synaptic pairing. Synapses from the right eye were identified by pairing VGluT2 clusters with CTB (Alexa 488) clusters. The volume of each cluster reflected the total voxel volume of all connected voxels, and the total signal intensity was a sum of voxel intensity within the volume of the connected voxels.

Revised version: To identify paired pre-and postsynaptic clusters, we first measured the centroid-centroid distance of each cluster in the Dy749P1 (Homer1) and Alexa 647 (Bassoon) channels to the closest cluster in the other channel. We next quantified the signal intensity of each opposing synaptic channel within a 140 nm shell surrounding each cluster. A 2D histogram was plotted based on the measured centroid-centroid distances and opposing channel signal densities of each cluster. Paired clusters with closely positioned centroids and high intensities of apposed channel signal were identified using the OPTICS algorithm. Retinogeniculate synapses were identified by pairing Bassoon (Alexa 647) clusters with VGluT2 (Cy3B) clusters using the same method as pre/post-synaptic pairing. Synapses from the right eye were identified by pairing VGluT2 clusters with CTB (Alexa 488) clusters. The volume of each cluster reflected the total voxel volume of all connected voxels, and the total signal intensity was a sum of voxel intensity within the volume of the connected voxels.

Multi-AZ synapse identification and quantification

Earlier version: To determine whether an eye-specific VGluT2 cluster is a mAZ synapse or a sAZ synapse, we measured the number of active zones (defined by individual Bassoon clusters) associated with each VGluT2 cluster in the dataset. A 3D shell was extended 140 nm from the surface voxels of each VGluT2 cluster and any Bassoon clusters that fell within the shell were considered to be associated with the target VGluT2 cluster. The number of active zones (AZs) associated with each VGluT2 cluster was then measured. VGluT2 clusters associated with more than 1 AZ were defined as mAZ synapses, while those associated with only 1 AZ were defined as sAZ synapses.

Revised version: To determine whether an eye-specific VGluT2 cluster is a mAZ synapse or a sAZ synapse, we measured the number of active zones (defined by individual Bassoon clusters) associated with each VGluT2 cluster in the dataset. A 3D shell was extended 140 nm from the surface voxels of each VGluT2 cluster and any Bassoon clusters that fell within the shell were considered to be associated with the target VGluT2 cluster.

Added in the revised version: The number of active zones (AZs) associated with each VGluT2 cluster was then measured. VGluT2 clusters associated with more than 1 AZ were defined as mAZ synapses, while those associated with only 1 AZ were defined as sAZ synapses.

Earlier version: Quantification of mAZ and sAZ synapse VGluT2 cluster volume was performed using the “regionprops” function in MATLAB, which provided the voxel size and weighted centroid of each VGluT2 cluster. The search for sAZ synapses adjacent to mAZ synapses (synaptic clustering analysis) was conducted using a similar search approach as for associated Bassoon clusters, with expansion shell sizes ranging from 1 μm to 4 μm from the surface voxels of each mAZ synapse. The main figures in the study utilized an expansion size of 1.5 μm. An eye-specific sAZ synapse was considered to be near a mAZ synapse if its weighted centroid fell within the expanded region.

Revised version: Quantification of mAZ and sAZ synapse VGluT2 cluster volume was performed using the “regionprops” function in MATLAB, which provided the voxel size and weighted centroid of each VGluT2 cluster. The search for sAZ synapses adjacent to mAZ synapses (synaptic clustering analysis) was conducted using a similar search approach as for associated Bassoon clusters, with expansion shell sizes ranging from 1 μm to 4 μm from the surface voxels of each mAZ synapse. The main Figs. in the study utilized an expansion distance of 1.5 μm. An eye-specific sAZ synapse was considered to be near a mAZ synapse if its weighted centroid fell within the expanded volume.

Quantification and statistical analysis

Added in the revised version: Statistical analysis was performed using SPSS. Plots were generated by SPSS or R (ggplot2). Statistical details are presented in the figure legends and supplemental tables S2/S3. For all measurements in this paper, we analyzed N = 3 biological replicates (individual mice) for each genotype (WT and β2KO) at each age (P2, P4, and P8).

Earlier version: Statistical analysis was performed using SPSS. Plots were generated by SPSS or R (ggplot2). Statistical details are presented in the figure legends. For all measurements in this paper, we analyzed N = 3 biological replicates (individual mice) for each genotype (WT and β2KO) at each age (P2, P4, and P8). Cluster densities, synapse AZ number, average VGluT2 cluster volume, and all fractional measurements were presented as mean ± SEM values in paired bar graphs and statistical analysis was performed using paired-T tests. Nonparametric Kolmogorov-Smirnov tests were used in all cumulative histogram comparisons. We used a linear mixed model to compare VGluT2 cluster volumes (Figure 2) and the distance measurements in Figure 4B. For VGluT2 cluster volume comparisons, the age or eye-of-origin was the fixed main factor and biological replicate IDs were nested random factors. In distance measurement comparisons, the mAZ synapse AZ number was the fixed main factor and biological replicate IDs were nested random factors. In violin plots, each violin showed the distribution of grouped data from all biological replicates from the same condition. Each black dot represents the median value of each biological replicate, and the horizontal black line represents the group median. Black lines connect measurements of CTB(+) and CTB(-) populations from the same biological replicate. Asterisks in all figures indicate statistical significance: *P<0.05, **P<0.01, ***P<0.001. For VGluT2 volume comparisons, we calculated 5/95% confidence intervals based on the mixed linear model. 5/95% confidence intervals in synapse densities, AZ numbers, and distances were calculated for paired or unpaired data comparisons using SPSS.

Revised version: Cluster densities, synapse AZ number, average VGluT2 cluster volume, and all fractional measurements were presented as mean ± SEM values in paired bar graphs and statistical analysis was performed using two-tailed paired-T tests. Nonparametric Kolmogorov-Smirnov tests were used in all cumulative histogram comparisons. We used a linear mixed model to compare VGluT2 cluster volumes (Figs. 3 and S2) and distance measurements (Figs. 5 and S5). For VGluT2 cluster volume comparisons, the age or eye-of-origin was the fixed main factor and biological replicate IDs were nested random factors. In distance measurement comparisons, the mAZ synapse AZ number was the fixed main factor and biological replicate IDs were nested random factors. In violin plots, each violin showed the distribution of grouped data from all biological replicates from the same condition. Each black dot represents the median value of each biological replicate, and the horizontal black line represents the group median. Black lines connect measurements of CTB(+) and CTB(-) populations from the same biological replicate. Asterisks in all figures indicate statistical significance: *P<0.05, **P<0.01, ***P<0.001. For VGluT2 volume comparisons, we calculated 5/95% confidence intervals based on the linear mixed model. 5/95% confidence intervals in synapse densities, AZ numbers, and distances were calculated for paired or unpaired data comparisons using SPSS.

Added in the revised version: All statistical results are listed in Table S2. All conclusions we draw in the main text from the data have corresponding confidence intervals listed in both the figure legends and Table S3.